<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 115431  -->
  <question type="stack">
    <name>
      <text>000.MCQ-rb-2026-04-04</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>[[lang code='ja']][[/lang]][[lang code='en']][[/lang]]<br></p>
<p>{@qtext@}</p>
<p>[[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Latest version is github/yositomi-opu/stack_question */
stack_include("https://stack.mathedu.jp/sc/ky_linear_algebra.txt");
/* Number of [all options, correct options] (rb/radio button case does work) */
[oN, cN]:[5, 2];
/* key parameters(e.g. which control the only level of the question) here if required */
[rw, cl]:rand([[3, 4], [4, 5], [3, 5]]);
pv:rand([[1, 2], [1, 3], [1, 2, 3]]);
/*********************************** FLAGS ***********************************/
/* Flags: modification not required if not necessary */
/* Flags for mcq_template(before include pre template) */
/* The difference of -cb and -rb is only type of ans1(checkbox or radio button) and this setting */
__mcq_rb_cb:"rb";
/* cN:(min or max)(cN, # of all correct patterns), wN:(min or max)(oN_cN, # of all incorrect patterns) */
__mcq_min_max:min;
/* display feedback? ( true recommended ) */
__mcq_fbp:true;
/* use subscore ? cf. ATCM2022 Yoshitomi */
__mcq_subsc:true;
/* include "I have no idea option? " (default: false) */
__mcq_noidea:false;
/* include "nocorrect answer" option ? (default: false) */
__mcq_nocorrectopt:false;
/* if include nocorrect answer options, is it correct answer or not ? */
__mcq_nocorrecttrue:rand([true, false]);
/* add any of "choose one" or "choose all" automatically to question text ? (default: true) */
__mcq_qtext_auto:true;
/****************************** DO NOT EDIT HERE *****************************/
stack_include("https://stack.mathedu.jp/sc/mcq_template_pre.txt");
/*****************************************************************************/

/**************** QUESTION VARIABLES for MAIN CONTENT ****************/
/* 正答・誤答パターン生成用変数, 問題文生成 */
matF:randechelon(rw, cl, pv, [-2, -1, -1, 0, 1, 1, 2]);
matP:random_matrix_det_list(rw, [-2, -1, -1, 0, 1, 1, 2], [-1, 1]);
matA:matP.matF;
/* qtext : 問題文の主文(Main sentence for question text/"選べ"は自動) */
if %_STACK_LANG="ja" then qtext:sconcat("\\(A=", tex1(matA), "\\)のとき, 行簡約形として正しいものを __SELTYPE__.") else qtext:sconcat("When \\(A=", tex1(matA), "\\), then choose __SELTYPE__ as the row reduced echelonform.");
if %_STACK_LANG="ja" then _Cm1:"正解パターン1です." else _Cm1:"Correct pattern 1";
if %_STACK_LANG="ja" then _Cm2:"正解パターン2です." else _Cm2:"Correct pattern 2";
if %_STACK_LANG="ja" then _Wm1:"誤答パターン1です." else _Wm1:"Incorrect pattern 1";
if %_STACK_LANG="ja" then _Wm2:"誤答パターン1です." else _Wm2:"Incorrect pattern 2";
if %_STACK_LANG="ja" then _Wm3:"誤答パターン1です." else _Wm3:"Incorrect pattern 3";
/************************ Patterns ********************/
/* 正解/誤答のパターンとフィードバックをセットで List に push していく. */
addCopt(_rp(["Correct pattern1-1"]), _Cm1);
addCopt(_rp(["Correct pattern2-1"]), _Cm2);
addWopt(_rp(["Wpattern1-1"]), _Wm1);
addWopt(_rp(["Wpattern2-1"]), _Wm2);
addWopt(_rp(["Wpattern3-1", "Wpattern3-2"]), _Wm3);
/* 各パターンからは1つ以上選択される. 何個ずつ選ぶかは自動で行う(mkPattern関数). */
/**************** END OF QUESTION VARIABLES for MAIN CONTENT ******************/

/****************************** DO NOT EDIT HERE *****************************/
stack_include("https://stack.mathedu.jp/sc/mcq_template_post.txt");
/*****************************************************************************/]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="html">
      <text>{@matA@}, {@matF@}, {@map(first,ta1)@}:{@ta1@}</text>
    </questionnote>
    <questiondescription format="html">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> 
<span lang="ja" class="multilang">よくできました。正解です！</span>
<span lang="en" class="multilang">Well done! Spot on!</span>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span>
<span lang="ja" class="multilang">部分的に正解です。</span>
<span lang="en" class="multilang">Partially correct!</span>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span>
<span lang="ja" class="multilang">間違いです。</span>
<span lang="en" class="multilang">Incorrect!</span>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>(</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>radio</type>
      <tans>ta1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered,hideanswer</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[/* ja:ラジオボタンの場合は, ans1 は数値， チェックボックスのときは, リスト */
/* en:radio button case: ans1 is numerical , checkbox case: ans1 is list */
if listp(ans1) then chk1:is(ans1=Cans1) else chk1:member(ans1,Cans1);
if listp(ans1) then ans1L:ans1 else ans1L:[ans1];
ans1WL:makelist(floor(ans1L[k]/_cw_delimiter),k,1,length(ans1L));
ans1CL:makelist(if ans1L[k]<_cw_delimiter then floor(ans1L[k]/_max_opts) else -1,k,1,length(ans1L));
countL(L,t):=block([k,rt:0],for k:1 thru length(L) do if L[k]=t then rt:rt+1,return(rt));
cAN:countL(ans1WL,0);
wAN:apply("+",makelist(countL(ans1WL,k),k,1,_max_patterns));
flgCL:makelist(member(k,ans1CL),k,1,length(CLbase));
flgWL:makelist(member(k,ans1WL),k,1,length(WLbase));
flgWnc:member(_nocorrect_flag,ans1WL);
flgC0:member(0,ans1CL);
flgNOID:member(_noidea_flag,ans1WL);
msg:"";
/* customize appropriate message for each incorrect(wrong) pattern */
for k:1 thru length(CLbase) do if flgCL[k] then msg:sconcat(msg,CmsgL[k],"<br>");
for k:1 thru length(WLbase) do if flgWL[k] then msg:sconcat(msg,WmsgL[k],"<br>");
if flgC0 then msg:sconcat(msg,"選択肢に正解はありません. 他を選択していなければ正解です.<br>");
if flgWnc then msg:sconcat(msg,"正解の選択肢はあります!<br>");
sc:cAN/(cN+wAN)*ssc;
/* 選択肢に正解がある(nocorrecttrue:false)のに, 正解がないを選択した場合(flg9) は, 常に 0 点 */
if (not nocorrecttrue) and flgWnc then sc:0;
/* 選択肢に正解がない時, 正解がないを選択しなかった場合は, cAN=0 なので自動的に 0点 */
/* 選択肢に正解がない時, 他の誤答選択肢を選択した場合は 0点 */
if nocorrecttrue and apply("or", flgWL) then sc:0;
if not fbp then msg:"";]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>NOIDEA-option-detection</description>
        <answertest>AlgEquiv</answertest>
        <sans>flgNOID</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<!-- 解答分析で PRT1-2T は, わからないを選択したことを検出 -->
[[lang code='ja']]「わからない」を検出(得点は0)[[/lang]]
[[lang code='en']] Detect "no idea options"(score is 0)[[/lang]]]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>採点とFeedbackは解答変数で設定</description>
        <answertest>AlgEquiv</answertest>
        <sans>chk1</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<!-- ラジオボタンで正解オプションが 2個以上ある場合に対応 -->
<p>
[[lang code='ja']]{@msg@}[[/lang]]
[[lang code='en']]{@msg@}[[/lang]]
</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>sc</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<!-- general error / 一般の誤り -->
<p>
[[lang code='ja']]{@msg@}[[/lang]]
[[lang code='en']]{@msg@}[[/lang]]
</p>
<p>
[[lang code='ja']]得点は{@sc@}です.[[/lang]]
[[lang code='en']]Your score is {@sc@}.[[/lang]]
</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>10</name>
        <description>Wrong(Incorrect) Pattern 6</description>
        <answertest>AlgEquiv</answertest>
        <sans>flgWL[6]</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-11-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-11-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Correct Pattern 1</description>
        <answertest>AlgEquiv</answertest>
        <sans>flgCL[1]</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<!-- dummy -->
<p>[[lang code="ja"]][[/lang]][[lang code="en"]][[/lang]]</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<!-- dummy -->
<p>[[lang code="ja"]][[/lang]][[lang code="en"]][[/lang]]</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Correct Pattern 2</description>
        <answertest>AlgEquiv</answertest>
        <sans>flgCL[2]</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>4</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<!-- dummy -->
<p>[[lang code="ja"]][[/lang]][[lang code="en"]][[/lang]]</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<!-- dummy -->
<p>[[lang code="ja"]][[/lang]][[lang code="en"]][[/lang]]</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>Correct Pattern 3</description>
        <answertest>AlgEquiv</answertest>
        <sans>flgCL[3]</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>5</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<!-- dummy -->
<p>[[lang code="ja"]][[/lang]][[lang code="en"]][[/lang]]</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>5</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<!-- dummy -->
<p>[[lang code="ja"]][[/lang]][[lang code="en"]][[/lang]]</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>5</name>
        <description>Wrong(Incorrect) pattern 1</description>
        <answertest>AlgEquiv</answertest>
        <sans>flgWL[1]</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>6</truenextnode>
        <trueanswernote>prt1-6-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<!-- dummy -->
<p>[[lang code="ja"]][[/lang]][[lang code="en"]][[/lang]]</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>6</falsenextnode>
        <falseanswernote>prt1-6-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<!-- dummy -->
<p>[[lang code="ja"]][[/lang]][[lang code="en"]][[/lang]]</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>6</name>
        <description>Wrong(Incorrect) Pattern 2</description>
        <answertest>AlgEquiv</answertest>
        <sans>flgWL[2]</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>7</truenextnode>
        <trueanswernote>prt1-7-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<!-- dummy -->
<p>[[lang code="ja"]][[/lang]][[lang code="en"]][[/lang]]</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>7</falsenextnode>
        <falseanswernote>prt1-7-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<!-- dummy -->
<p>[[lang code="ja"]][[/lang]][[lang code="en"]][[/lang]]</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>7</name>
        <description>Wrong(Incorrect) Pattern 3</description>
        <answertest>AlgEquiv</answertest>
        <sans>flgWL[3]</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>8</truenextnode>
        <trueanswernote>prt1-8-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<!-- dummy -->
<p>[[lang code="ja"]][[/lang]][[lang code="en"]][[/lang]]</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>8</falsenextnode>
        <falseanswernote>prt1-8-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<!-- dummy -->
<p>[[lang code="ja"]][[/lang]][[lang code="en"]][[/lang]]</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>8</name>
        <description>Wrong(Incorrect) Pattern 4</description>
        <answertest>AlgEquiv</answertest>
        <sans>flgWL[4]</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>9</truenextnode>
        <trueanswernote>prt1-9-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<!-- dummy -->
<p>[[lang code="ja"]][[/lang]][[lang code="en"]][[/lang]]</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>9</falsenextnode>
        <falseanswernote>prt1-9-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<!-- dummy -->
<p>[[lang code="ja"]][[/lang]][[lang code="en"]][[/lang]]</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>9</name>
        <description>Wrong(Incorrect) Pattern 5</description>
        <answertest>AlgEquiv</answertest>
        <sans>flgWL[5]</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>10</truenextnode>
        <trueanswernote>prt1-10-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>10</falsenextnode>
        <falseanswernote>prt1-10-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>