<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 109912  -->
  <question type="stack">
    <name>
      <text>000.MCQ-rb-2026-01-27</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p><br></p>
<p>{@qtext@}</p>
<p>[[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[stack_include("https://stack.mathedu.jp/sc/ky_linear_algebra.txt");
/* [全選択肢の数, 正解選択肢の数] (rb のときも2個以上可能) */
[oN, cN]:[10, 3];
/* random parameters (設定を変えて問題を複製するときに変える or 固定する部分) */
[rw, cl]:rand([[3, 4], [4, 5], [3, 5]]);
pv:rand([[1, 2], [1, 3], [1, 2, 3]]);
/*********************************** FLAGS ***********************************/
/* 以下は不要ならさわらなくてよい(modification not required if not necessary) */
/* Flags for mcq_template(before include pre template) */
/* -cb と -rb は解答変数の型(checkbox or radio button) と以下の設定の違いのみ */
__mcq_rb_cb:"rb";
/* cN:(min or max)(cN, 正答パターン数), wN:(min or max)(oN_cN, 誤答パターン数) */
__mcq_min_max:min;
/* feedback を表示するか? (display feedback ? ) */
__mcq_fbp:true;
/* 部分点をつけるか? (use subscore ? cf. ATCM2022 ) */
__mcq_subsc:true;
/* "わからない"選択肢を含めるか? (false) */
__mcq_noidea:false;
/* "選択肢に正解はない"選択肢を含めるか? (false) */
__mcq_nocorrectopt:false;
/* 選択肢に正解はない選択肢は正解か? (__mcq_nocorrectopt と自動で and をとる) */
__mcq_nocorrecttrue:rand([true, false]);
 /* 自動で"1つ選べ", "すべて選べ" を付加するか? (true) */
__mcq_qtext_auto:true;
/****************************** DO NOT EDIT HERE *****************************/
stack_include("https://stack.mathedu.jp/sc/mcq_template_pre.txt");
/*****************************************************************************/

/**************** QUESTION VARIABLES for MAIN CONTENT ****************/
/* 正答・誤答パターン生成用変数, 問題文生成 */
matF:randechelon(rw, cl, pv, [-2, -1, -1, 0, 1, 1, 2]);
matP:random_matrix_det_list(rw, [-2, -1, -1, 0, 1, 1, 2], [-1, 1]);
matA:matP.matF;
/* qtext : 問題文の主文(Main sentence for question text/"選べ"は自動) */
qtext:sconcat("\\(A=\\)", tex2(matA), "のとき, 行簡約形として正しいものを");
/************************ Patterns ********************/
/* 正解/誤答のパターンとフィードバックをセットで List に push していく. */
addCopt(_rp(["正解パターン1-1"]), "正答パターン1を選択しています. そのような考え方もありますね!");
addCopt(_rp(["正解パターン2-1"]), "正答パターン2を選択しています. そのような考え方もありますね!");
addCopt(_rp(["正解パターン3-1","正解パターン3-2"]), "標準的ではないですが, 正解です.");
addWopt(_rp(["誤答パターン1-1"]), "Feedback message for wrong pattern 1");
addWopt(_rp(["誤答パターン2-1"]), "Feedback message for wrong pattern 2");
addWopt(_rp(["誤答パターン3-1","誤答パターン3-2"]), "Feedback message for wrong pattern 3");
addWopt(_rp(["誤答パターン4-1","誤答パターン4-2"]), "Feedback message for wrong pattern 4");
/* 各パターンからは1つ以上選択される. 何個ずつ選ぶかは自動で行う(mkPattern関数). */
/**************** END OF QUESTION VARIABLES for MAIN CONTENT ******************/

/****************************** DO NOT EDIT HERE *****************************/
stack_include("https://stack.mathedu.jp/sc/mcq_template_post.txt");
/*****************************************************************************/]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="html">
      <text>{@matA@}, {@matF@}, {@map(first,ta1)@}:{@ta1@}</text>
    </questionnote>
    <questiondescription format="html">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> 
<span lang="ja" class="multilang">よくできました。正解です！</span>
<span lang="en" class="multilang">Well done! Spot on!</span>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span>
<span lang="ja" class="multilang">部分的に正解です。</span>
<span lang="en" class="multilang">Partially correct!</span>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span>
<span lang="ja" class="multilang">間違いです。</span>
<span lang="en" class="multilang">Incorrect!</span>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>(</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>radio</type>
      <tans>ta1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered,hideanswer</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[/* ja:ラジオボタンの場合は, ans1 は数値， チェックボックスのときは, リスト */
/* en:radio button case: ans1 is numerical , checkbox case: ans1 is list */
if listp(ans1) then chk1:is(ans1=Cans1) else chk1:member(ans1,Cans1);
if listp(ans1) then ans1L:ans1 else ans1L:[ans1];
ans1WL:makelist(floor(ans1L[k]/_cw_delimiter),k,1,length(ans1L));
ans1CL:makelist(if ans1L[k]<_cw_delimiter then floor(ans1L[k]/_max_opts) else -1,k,1,length(ans1L));
countL(L,t):=block([k,rt:0],for k:1 thru length(L) do if L[k]=t then rt:rt+1,return(rt));
cAN:countL(ans1WL,0);
wAN:apply("+",makelist(countL(ans1WL,k),k,1,_max_patterns));
flgCL:makelist(member(k,ans1CL),k,1,length(CLbase));
flgWL:makelist(member(k,ans1WL),k,1,length(WLbase));
flgWnc:member(_nocorrect_flag,ans1WL);
flgC0:member(0,ans1CL);
flgNOID:member(_noidea_flag,ans1WL);
msg:"";
/* customize appropriate message for each incorrect(wrong) pattern */
for k:1 thru length(CLbase) do if flgCL[k] then msg:sconcat(msg,CmsgL[k],"<br>");
for k:1 thru length(WLbase) do if flgWL[k] then msg:sconcat(msg,WmsgL[k],"<br>");
if flgC0 then msg:sconcat(msg,"選択肢に正解はありません. 他を選択していなければ正解です.<br>");
if flgWnc then msg:sconcat(msg,"正解の選択肢はあります!<br>");
sc:cAN/(cN+wAN)*ssc;
/* 選択肢に正解がある(nocorrecttrue:false)のに, 正解がないを選択した場合(flg9) は, 常に 0 点 */
if (not nocorrecttrue) and flgWnc then sc:0;
/* 選択肢に正解がない時, 正解がないを選択しなかった場合は, cAN=0 なので自動的に 0点 */
/* 選択肢に正解がない時, 他の誤答選択肢を選択した場合は 0点 */
if nocorrecttrue and apply("or", flgWL) then sc:0;
if not fbp then msg:"";]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>NOIDEA-option-detection</description>
        <answertest>AlgEquiv</answertest>
        <sans>flgNOID</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<!-- 解答分析で PRT1-2T は, わからないを選択したことを検出 -->
「わからない」を検出 (Detect "no idea options") \(\to\) score :0 and do not check the answer.]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>採点とFeedbackは解答変数で設定</description>
        <answertest>AlgEquiv</answertest>
        <sans>chk1</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<!-- ラジオボタンで正解オプションが 2個以上ある場合に対応 -->
<p>{@msg@}</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>sc</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<!-- general error / 一般の誤り -->
<p>{@msg@}</p>
<p>得点は{@sc@}です.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1182809594</deployedseed>
    <deployedseed>2085548457</deployedseed>
    <deployedseed>773701273</deployedseed>
    <deployedseed>655472694</deployedseed>
    <deployedseed>1249613553</deployedseed>
    <deployedseed>519077920</deployedseed>
    <deployedseed>1738289652</deployedseed>
    <deployedseed>131970873</deployedseed>
    <deployedseed>1488805594</deployedseed>
    <deployedseed>2037850386</deployedseed>
    <deployedseed>862677479</deployedseed>
    <deployedseed>1840365972</deployedseed>
    <deployedseed>314207632</deployedseed>
    <deployedseed>1110743210</deployedseed>
    <deployedseed>642613939</deployedseed>
    <deployedseed>671972308</deployedseed>
    <deployedseed>897902351</deployedseed>
    <deployedseed>1684651903</deployedseed>
    <deployedseed>1516829949</deployedseed>
    <deployedseed>1494872326</deployedseed>
    <deployedseed>1380356300</deployedseed>
    <deployedseed>2081045341</deployedseed>
    <deployedseed>727489884</deployedseed>
    <deployedseed>1115536029</deployedseed>
    <deployedseed>372101839</deployedseed>
    <deployedseed>1541654861</deployedseed>
    <deployedseed>1032707530</deployedseed>
    <deployedseed>867940411</deployedseed>
    <deployedseed>4949044</deployedseed>
    <deployedseed>1830033960</deployedseed>
    <deployedseed>876312283</deployedseed>
    <deployedseed>1248258634</deployedseed>
    <deployedseed>1306075324</deployedseed>
    <deployedseed>66783509</deployedseed>
    <deployedseed>1989922073</deployedseed>
    <deployedseed>1358643795</deployedseed>
    <deployedseed>1198521804</deployedseed>
    <deployedseed>128946728</deployedseed>
    <deployedseed>35526900</deployedseed>
    <deployedseed>412383118</deployedseed>
    <deployedseed>1335024755</deployedseed>
    <deployedseed>154575501</deployedseed>
    <deployedseed>1103934808</deployedseed>
    <deployedseed>1601955648</deployedseed>
    <deployedseed>1319264748</deployedseed>
    <deployedseed>1180333732</deployedseed>
    <deployedseed>6716332</deployedseed>
    <deployedseed>2049403979</deployedseed>
    <deployedseed>1316545688</deployedseed>
    <deployedseed>952831783</deployedseed>
    <deployedseed>1019007008</deployedseed>
    <deployedseed>1168377454</deployedseed>
    <deployedseed>654099589</deployedseed>
    <deployedseed>230255982</deployedseed>
    <deployedseed>1729431810</deployedseed>
    <deployedseed>574620488</deployedseed>
    <deployedseed>1880294916</deployedseed>
    <deployedseed>2018390154</deployedseed>
    <deployedseed>1050832529</deployedseed>
    <deployedseed>536029862</deployedseed>
    <deployedseed>762752135</deployedseed>
    <deployedseed>1415487858</deployedseed>
    <deployedseed>333536070</deployedseed>
    <deployedseed>1416438983</deployedseed>
    <deployedseed>925357110</deployedseed>
    <deployedseed>1391184900</deployedseed>
    <deployedseed>1805050280</deployedseed>
    <deployedseed>1277076329</deployedseed>
    <deployedseed>284093830</deployedseed>
    <deployedseed>1580465621</deployedseed>
    <deployedseed>745010445</deployedseed>
    <deployedseed>897369620</deployedseed>
    <deployedseed>1946800022</deployedseed>
    <deployedseed>315498541</deployedseed>
    <deployedseed>1631330351</deployedseed>
    <deployedseed>288470029</deployedseed>
    <deployedseed>2081239903</deployedseed>
    <deployedseed>492249310</deployedseed>
    <deployedseed>564437228</deployedseed>
    <deployedseed>282808718</deployedseed>
    <deployedseed>881437182</deployedseed>
    <deployedseed>850037971</deployedseed>
    <deployedseed>780794114</deployedseed>
    <deployedseed>899848278</deployedseed>
    <deployedseed>38588674</deployedseed>
    <deployedseed>1023671254</deployedseed>
    <deployedseed>74587293</deployedseed>
    <deployedseed>1811475422</deployedseed>
    <deployedseed>1185622847</deployedseed>
    <deployedseed>163687487</deployedseed>
    <deployedseed>1002956915</deployedseed>
    <deployedseed>530169365</deployedseed>
    <deployedseed>1746535065</deployedseed>
    <deployedseed>540551071</deployedseed>
    <deployedseed>1603893496</deployedseed>
    <deployedseed>598325379</deployedseed>
    <deployedseed>1180706914</deployedseed>
    <deployedseed>184794274</deployedseed>
    <deployedseed>1878184124</deployedseed>
    <deployedseed>1189920308</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>管理者のインプットを想定したテストケースは満点です。</description>
      <testinput>
        <name>ans1</name>
        <value>mcq_correct(ta1)[1]</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>