<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 115488  -->
  <question type="stack">
    <name>
      <text>001.GaussElimRowReducedCheck-B-3x4rk2-rb</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[[[lang code="en"]][[/lang]][[lang code="ja"]][[/lang]][[lang code="fr"]][[/lang]][[lang code="it"]][[/lang]][[lang code="de"]][[/lang]][[lang code="pt"]][[/lang]][[lang code="zh"]][[/lang]][[lang code="ko"]][[/lang]][[lang code="ru"]][[/lang]][[lang code="sv"]][[/lang]][[lang code="es"]][[/lang]]{@%__mcq_langcode@}
<p>{@qtext@}</p>
<p>[[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/************** You can find latest version in github/yositomi-opu/stack_question ***********/
stack_include("https://yositomi-opu.github.io/stack_questions/ky_linear_algebra.mac");
/* The difference of -cb and -rb is only type of ans1(checkbox or radio button) and this setting */
%__mcq_rb_cb:"rb";
/*************** some usual fuctions I use. Feel free to customize here **********************/
powerdisp:true;
%_rp(L):=random_permutation(L);
stack_include("https://yositomi-opu.github.io/stack_questions/tex_library.mac");
/*********************************** Flags ***********************************/
/* NOTE: opts=options, copts=correct options, wopts=wrong(incorrect) options */

/* Flags : modification not required if not necessary. See mcq_template_pre_cas.txt for details */
/* [use feedback?, use subscore?, use noidea opt, use nocorrect opt, is nocorrect opt true?, use qtext_auto?] */

%_MCQ_FLAGS:[true, true, false, false, rand([true, false]), true];

/* NOTE: If the number of correct/incorrect answer patterns exceeds 5/9, increase the number of PRT nodes */
/*       and adjust the following constants. */
%__mcq_max_cp:5;
%__mcq_max_wp:9;

/********************************* DO NOT EDIT HERE **************************************/
stack_include("https://yositomi-opu.github.io/stack_questions/mcq_template_pre_cas.mac");
/*****************************************************************************************/
/* MCQ_WEBAPP_DATA_BASE64:eyJ2ZXJzaW9uIjoxLCJxdWVzdGlvbklkIjoiR2F1c3NFbGltUm93UmVkdWNlZENoZWNrLUItM3g0cmsyIiwibW9kZSI6InJiIiwiYmFzZUxhbmd1YWdlIjoiamEiLCJhY3RpdmVMYW5ndWFnZXMiOlsiZW4iLCJqYSIsImZyIiwiaXQiLCJkZSIsInB0IiwiemgiLCJrbyIsInJ1Iiwic3YiLCJlcyJdLCJxdWVzdGlvbkxhbmd1YWdlSW5kZXBlbmRlbnQiOmZhbHNlLCJxdWVzdGlvbnMiOnsiZW4iOnsidmFsdWUiOiJBcHBseWluZyBlbGVtZW50YXJ5IHJvdyBvcGVyYXRpb25zIHRvIHRoZSB7QCVfbnJAfVxcdGltZXN7JV9uY30gbWF0cml4IFxcKEE9XFwpIHtAbWF0QXBAfSB5aWVsZHMge0BtYXRGcUB9LiBTZWxlY3QgdGhlIGNvcnJlY3QgcGFpciBvZiB2YWx1ZXMgZm9yIFxcKHAsIHFcXCkuIiwidHlwZSI6ImNhc3RleHQifSwiamEiOnsidmFsdWUiOiJ7QCVfbnJAfVxcdGltZXN7JV9uY30g6KGM5YiXXFwoQT1cXCkge0BtYXRBcEB9IOOCkuihjOOBq+mWouOBl+OBpuWfuuacrOWkieW9ouOBmeOCi+OBqCwge0BtYXRGcUB9IOOBq+OBquOCiy4g44GT44Gu44Go44GNLCBcXChwLCBxXFwp44Gu5YCk44Gu57WE44Go44GX44Gm5q2j44GX44GE44KC44Gu44KS6YG444G577yOIiwidHlwZSI6ImNhc3RleHQifSwiZnIiOnsidmFsdWUiOiJMJ2FwcGxpY2F0aW9uIGQnb3DDqXJhdGlvbnMgw6lsw6ltZW50YWlyZXMgc3VyIGxlcyBsaWduZXMgZGUgbGEgbWF0cmljZSB7QCVfbnJAfVxcdGltZXN7JV9uY30gXFwoQT1cXCkge0BtYXRBcEB9IGRvbm5lIHtAbWF0RnFAfS4gU8OpbGVjdGlvbm5leiBsZSBjb3VwbGUgY29ycmVjdCBkZSB2YWxldXJzIGRlIFxcKHAsIHFcXCkuIiwidHlwZSI6ImNhc3RleHQifSwiaXQiOnsidmFsdWUiOiJBcHBsaWNhbmRvIG9wZXJhemlvbmkgZWxlbWVudGFyaSBzdWxsZSByaWdoZSBhbGxhIG1hdHJpY2Uge0AlX25yQH1cXHRpbWVzeyVfbmN9IFxcKEE9XFwpIHtAbWF0QXBAfSwgc2kgb3R0aWVuZSB7QG1hdEZxQH0uIFNlbGV6aW9uYSBsYSBjb3BwaWEgY29ycmV0dGEgZGkgdmFsb3JpIGRpIFxcKHAsIHFcXCkuIiwidHlwZSI6ImNhc3RleHQifSwiZGUiOnsidmFsdWUiOiJEdXJjaCBlbGVtZW50YXJlIFplaWxlbnVtZm9ybXVuZ2VuIHdpcmQgZGllIHtAJV9uckB9XFx0aW1lc3slX25jfS1NYXRyaXggXFwoQT1cXCkge0BtYXRBcEB9IGluIHtAbWF0RnFAfSDDvGJlcmbDvGhydC4gV8OkaGxlbiBTaWUgZGFzIHJpY2h0aWdlIFdlcnRlcGFhciBmw7xyIFxcKHAsIHFcXCkgYXVzLiIsInR5cGUiOiJjYXN0ZXh0In0sInB0Ijp7InZhbHVlIjoiQXBsaWNhbmRvIG9wZXJhw6fDtWVzIGVsZW1lbnRhcmVzIHNvYnJlIGFzIGxpbmhhcyDDoCBtYXRyaXoge0AlX25yQH1cXHRpbWVzeyVfbmN9IFxcKEE9XFwpIHtAbWF0QXBAfSwgb2J0w6ltLXNlIHtAbWF0RnFAfS4gU2VsZWNpb25lIG8gcGFyIGNvcnJldG8gZGUgdmFsb3JlcyBkZSBcXChwLCBxXFwpLiIsInR5cGUiOiJjYXN0ZXh0In0sInpoIjp7InZhbHVlIjoi5a+5e0AlX25yQH1cXHRpbWVzeyVfbmN955+p6Zi1XFwoQT1cXCkge0BtYXRBcEB96L+b6KGM5Yid562J6KGM5Y+Y5o2i5ZCO77yM5b6X5Yiwe0BtYXRGcUB944CC6K+36YCJ5oupXFwocCwgcVxcKeeahOato+ehruWPluWAvOe7hOWQiOOAgiIsInR5cGUiOiJjYXN0ZXh0In0sImtvIjp7InZhbHVlIjoie0AlX25yQH1cXHRpbWVzeyVfbmN9IO2WieugrCBcXChBPVxcKSB7QG1hdEFwQH3sl5Ag6riw67O4IO2WiSDsl7DsgrDsnYQg7KCB7Jqp7ZWY66m0IHtAbWF0RnFAfeqwgCDrkJzri6QuIOydtOuVjCBcXChwLCBxXFwp7J2YIOqwkuycvOuhnCDsmKzrsJTrpbgg7Iic7ISc7IyN7J2EIOqzoOultOyLnOyYpC4iLCJ0eXBlIjoiY2FzdGV4dCJ9LCJydSI6eyJ2YWx1ZSI6ItCf0L7RgdC70LUg0Y3Qu9C10LzQtdC90YLQsNGA0L3Ri9GFINC/0YDQtdC+0LHRgNCw0LfQvtCy0LDQvdC40Lkg0YHRgtGA0L7QuiDQvNCw0YLRgNC40YbQsCDRgNCw0LfQvNC10YDQsCB7QCVfbnJAfVxcdGltZXN7JV9uY30gXFwoQT1cXCkge0BtYXRBcEB9INC/0YDQuNC90LjQvNCw0LXRgiDQstC40LQge0BtYXRGcUB9LiDQktGL0LHQtdGA0LjRgtC1INC/0YDQsNCy0LjQu9GM0L3Rg9GOINC/0LDRgNGDINC30L3QsNGH0LXQvdC40LkgXFwocCwgcVxcKS4iLCJ0eXBlIjoiY2FzdGV4dCJ9LCJzdiI6eyJ2YWx1ZSI6Ik7DpHIgZWxlbWVudMOkcmEgcmFkb3BlcmF0aW9uZXIgdXRmw7ZycyBww6Uge0AlX25yQH1cXHRpbWVzeyVfbmN9LW1hdHJpc2VuIFxcKEE9XFwpIHtAbWF0QXBAfSwgZXJow6VsbHMge0BtYXRGcUB9LiBWw6RsaiBkZXQga29ycmVrdGEgdsOkcmRlcGFyZXQgZsO2ciBcXChwLCBxXFwpLiIsInR5cGUiOiJjYXN0ZXh0In0sImVzIjp7InZhbHVlIjoiQWwgYXBsaWNhciBvcGVyYWNpb25lcyBlbGVtZW50YWxlcyBwb3IgZmlsYXMgYSBsYSBtYXRyaXoge0AlX25yQH1cXHRpbWVzeyVfbmN9IFxcKEE9XFwpIHtAbWF0QXBAfSwgc2Ugb2J0aWVuZSB7QG1hdEZxQH0uIFNlbGVjY2lvbmEgZWwgcGFyIGNvcnJlY3RvIGRlIHZhbG9yZXMgZGUgXFwocCwgcVxcKS4iLCJ0eXBlIjoiY2FzdGV4dCJ9fSwicm93cyI6W3sicGF0dGVybiI6IjAxIiwidHJ1dGgiOiJDIiwiY2hvaWNlX2VuIjoiTGlzdEFMMSIsImNob2ljZV90eXBlX2VuIjoiY2FzIiwiY2hvaWNlX2xpc3RfZXhwcl9lbiI6dHJ1ZSwiY2hvaWNlX2xhbmd1YWdlX2luZGVwZW5kZW50Ijp0cnVlLCJjaG9pY2VfZnIiOiJMaXN0QUwxIiwiY2hvaWNlX3R5cGVfZnIiOiJjYXMiLCJjaG9pY2VfbGlzdF9leHByX2ZyIjp0cnVlLCJjaG9pY2VfZGUiOiJMaXN0QUwxIiwiY2hvaWNlX3R5cGVfZGUiOiJjYXMiLCJjaG9pY2VfbGlzdF9leHByX2RlIjp0cnVlLCJjaG9pY2VfaXQiOiJMaXN0QUwxIiwiY2hvaWNlX3R5cGVfaXQiOiJjYXMiLCJjaG9pY2VfbGlzdF9leHByX2l0Ijp0cnVlLCJjaG9pY2VfamEiOiJMaXN0QUwxIiwiY2hvaWNlX3R5cGVfamEiOiJjYXMiLCJjaG9pY2VfbGlzdF9leHByX2phIjp0cnVlLCJmZWVkYmFja19lbiI6IkNvcnJlY3QuIFlvdSBzZWVtIHRvIHVuZGVyc3RhbmQgdGhlIHByb3BlcnRpZXMgb2YgcmVkdWNlZCByb3cgZWNoZWxvbiBtYXRyaWNlcy4iLCJmZWVkYmFja190eXBlX2VuIjoidGV4dCIsImZlZWRiYWNrX2ZyIjoiQ29ycmVjdC4gVm91cyBzZW1ibGV6IGNvbXByZW5kcmUgbGVzIHByb3ByacOpdMOpcyBkZXMgbWF0cmljZXMgc291cyBmb3JtZSDDqWNoZWxvbm7DqWUgcsOpZHVpdGUuIiwiZmVlZGJhY2tfdHlwZV9mciI6InRleHQiLCJmZWVkYmFja19kZSI6IlJpY2h0aWcuIFNpZSBzY2hlaW5lbiBkaWUgRWlnZW5zY2hhZnRlbiB2b24gTWF0cml6ZW4gaW4gcmVkdXppZXJ0ZXIgWmVpbGVuc3R1ZmVuZm9ybSB2ZXJzdGFuZGVuIHp1IGhhYmVuLiIsImZlZWRiYWNrX3R5cGVfZGUiOiJ0ZXh0IiwiZmVlZGJhY2tfaXQiOiJDb3JyZXR0by4gU2VtYnJhIGNoZSB0dSBhYmJpYSBjb21wcmVzbyBsZSBwcm9wcmlldMOgIGRlbGxlIG1hdHJpY2kgaW4gZm9ybWEgYSBzY2FsYSByaWRvdHRhIHBlciByaWdoZS4iLCJmZWVkYmFja190eXBlX2l0IjoidGV4dCIsImZlZWRiYWNrX2phIjoi5q2j6Kej44Gn44GZLiDooqvntITpmo7mrrXooYzliJco6KGM57Ch57SE6KGM5YiXKeOBruaAp+izquOCkueQhuino+OBl+OBpuOBhOOCi+OCiOOBhuOBp+OBmS4iLCJmZWVkYmFja190eXBlX2phIjoidGV4dCIsImZlZWRiYWNrX2J5X3RydXRoIjp0cnVlLCJjaG9pY2VfcHQiOiJMaXN0QUwxIiwiY2hvaWNlX3R5cGVfcHQiOiJjYXMiLCJjaG9pY2VfbGlzdF9leHByX3B0Ijp0cnVlLCJjaG9pY2VfemgiOiJMaXN0QUwxIiwiY2hvaWNlX3R5cGVfemgiOiJjYXMiLCJjaG9pY2VfbGlzdF9leHByX3poIjp0cnVlLCJjaG9pY2Vfa28iOiJMaXN0QUwxIiwiY2hvaWNlX3R5cGVfa28iOiJjYXMiLCJjaG9pY2VfbGlzdF9leHByX2tvIjp0cnVlLCJjaG9pY2VfcnUiOiJMaXN0QUwxIiwiY2hvaWNlX3R5cGVfcnUiOiJjYXMiLCJjaG9pY2VfbGlzdF9leHByX3J1Ijp0cnVlLCJjaG9pY2Vfc3YiOiJMaXN0QUwxIiwiY2hvaWNlX3R5cGVfc3YiOiJjYXMiLCJjaG9pY2VfbGlzdF9leHByX3N2Ijp0cnVlLCJjaG9pY2VfZXMiOiJMaXN0QUwxIiwiY2hvaWNlX3R5cGVfZXMiOiJjYXMiLCJjaG9pY2VfbGlzdF9leHByX2VzIjp0cnVlLCJmZWVkYmFja19wdCI6IkNvcnJldG8uIFBhcmVjZSBxdWUgY29tcHJlZW5kZSBhcyBwcm9wcmllZGFkZXMgZGFzIG1hdHJpemVzIG5hIGZvcm1hIGVzY2Fsb25hZGEgcmVkdXppZGEgcG9yIGxpbmhhcy4iLCJmZWVkYmFja190eXBlX3B0IjoidGV4dCIsImZlZWRiYWNrX3poIjoi5Zue562U5q2j56Gu44CC55yL5p2l5L2g5bey57uP55CG6Kej5LqG566A5YyW6KGM6Zi25qKv5b2i55+p6Zi155qE5oCn6LSo44CCIiwiZmVlZGJhY2tfdHlwZV96aCI6InRleHQiLCJmZWVkYmFja19rbyI6IuygleuLteyeheuLiOuLpC4g6riw7JW9IO2WiSDsgqzri6TrpqzqvLQg7ZaJ66Cs7J2YIOyEseyniOydhCDsnpgg7J207ZW07ZWY6rOgIOyeiOuKlCDqsoMg6rCZ7Iq164uI64ukLiIsImZlZWRiYWNrX3R5cGVfa28iOiJ0ZXh0IiwiZmVlZGJhY2tfcnUiOiLQktC10YDQvdC+LiDQn9C+LdCy0LjQtNC40LzQvtC80YMsINCy0Ysg0L/QvtC90LjQvNCw0LXRgtC1INGB0LLQvtC50YHRgtCy0LAg0LzQsNGC0YDQuNGGINCyINC/0YDQuNCy0LXQtNGR0L3QvdC+0Lkg0YHRgtGD0L/QtdC90YfQsNGC0L7QuSDRhNC+0YDQvNC1LiIsImZlZWRiYWNrX3R5cGVfcnUiOiJ0ZXh0IiwiZmVlZGJhY2tfc3YiOiJSw6R0dC4gRHUgdmVya2FyIGbDtnJzdMOlIGVnZW5za2FwZXJuYSBob3MgbWF0cmlzZXIgcMOlIHJlZHVjZXJhZCB0cmFwcHN0ZWdzZm9ybS4iLCJmZWVkYmFja190eXBlX3N2IjoidGV4dCIsImZlZWRiYWNrX2VzIjoiQ29ycmVjdG8uIFBhcmVjZSBxdWUgY29tcHJlbmRlcyBsYXMgcHJvcGllZGFkZXMgZGUgbGFzIG1hdHJpY2VzIGVuIGZvcm1hIGVzY2Fsb25hZGEgcmVkdWNpZGEgcG9yIGZpbGFzLiIsImZlZWRiYWNrX3R5cGVfZXMiOiJ0ZXh0In0seyJwYXR0ZXJuIjoiMDEiLCJ0cnV0aCI6IlciLCJjaG9pY2VfZW4iOiJMaXN0QkwxIiwiY2hvaWNlX3R5cGVfZW4iOiJjYXMiLCJjaG9pY2VfbGlzdF9leHByX2VuIjp0cnVlLCJjaG9pY2VfbGFuZ3VhZ2VfaW5kZXBlbmRlbnQiOnRydWUsImNob2ljZV9mciI6Ikxpc3RCTDEiLCJjaG9pY2VfdHlwZV9mciI6ImNhcyIsImNob2ljZV9saXN0X2V4cHJfZnIiOnRydWUsImNob2ljZV9kZSI6Ikxpc3RCTDEiLCJjaG9pY2VfdHlwZV9kZSI6ImNhcyIsImNob2ljZV9saXN0X2V4cHJfZGUiOnRydWUsImNob2ljZV9pdCI6Ikxpc3RCTDEiLCJjaG9pY2VfdHlwZV9pdCI6ImNhcyIsImNob2ljZV9saXN0X2V4cHJfaXQiOnRydWUsImNob2ljZV9qYSI6Ikxpc3RCTDEiLCJjaG9pY2VfdHlwZV9qYSI6ImNhcyIsImNob2ljZV9saXN0X2V4cHJfamEiOnRydWUsImZlZWRiYWNrX2VuIjoiT25lIG9mIHRoZSB2YWx1ZXMgaXMgY29ycmVjdCwgYnV0IHRoZSBvdGhlciBhcHBlYXJzIHRvIGJlIGluY29ycmVjdC4iLCJmZWVkYmFja190eXBlX2VuIjoidGV4dCIsImZlZWRiYWNrX2ZyIjoiTCd1bmUgZGVzIHZhbGV1cnMgZXN0IGNvcnJlY3RlLCBtYWlzIGwnYXV0cmUgc2VtYmxlIGluY29ycmVjdGUuIiwiZmVlZGJhY2tfdHlwZV9mciI6InRleHQiLCJmZWVkYmFja19kZSI6IkVpbmVyIGRlciBXZXJ0ZSBpc3QgcmljaHRpZywgYWJlciBkZXIgYW5kZXJlIHNjaGVpbnQgZmFsc2NoIHp1IHNlaW4uIiwiZmVlZGJhY2tfdHlwZV9kZSI6InRleHQiLCJmZWVkYmFja19pdCI6IlVubyBkZWkgdmFsb3JpIMOoIGNvcnJldHRvLCBtYSBsJ2FsdHJvIHNlbWJyYSBlc3NlcmUgZXJyYXRvLiIsImZlZWRiYWNrX3R5cGVfaXQiOiJ0ZXh0IiwiZmVlZGJhY2tfamEiOiLkuIDmlrnjga7lgKTjga/mraPjgZfjgYTjga7jgafjgZnjgYwsIOOCguOBhuS4gOaWueOBjOmBleOBhuOCiOOBhuOBp+OBmS4iLCJmZWVkYmFja190eXBlX2phIjoidGV4dCIsImZlZWRiYWNrX2J5X3RydXRoIjp0cnVlLCJjaG9pY2VfcHQiOiJMaXN0QkwxIiwiY2hvaWNlX3R5cGVfcHQiOiJjYXMiLCJjaG9pY2VfbGlzdF9leHByX3B0Ijp0cnVlLCJjaG9pY2VfemgiOiJMaXN0QkwxIiwiY2hvaWNlX3R5cGVfemgiOiJjYXMiLCJjaG9pY2VfbGlzdF9leHByX3poIjp0cnVlLCJjaG9pY2Vfa28iOiJMaXN0QkwxIiwiY2hvaWNlX3R5cGVfa28iOiJjYXMiLCJjaG9pY2VfbGlzdF9leHByX2tvIjp0cnVlLCJjaG9pY2VfcnUiOiJMaXN0QkwxIiwiY2hvaWNlX3R5cGVfcnUiOiJjYXMiLCJjaG9pY2VfbGlzdF9leHByX3J1Ijp0cnVlLCJjaG9pY2Vfc3YiOiJMaXN0QkwxIiwiY2hvaWNlX3R5cGVfc3YiOiJjYXMiLCJjaG9pY2VfbGlzdF9leHByX3N2Ijp0cnVlLCJjaG9pY2VfZXMiOiJMaXN0QkwxIiwiY2hvaWNlX3R5cGVfZXMiOiJjYXMiLCJjaG9pY2VfbGlzdF9leHByX2VzIjp0cnVlLCJmZWVkYmFja19wdCI6IlVtIGRvcyB2YWxvcmVzIGVzdMOhIGNvcnJldG8sIG1hcyBvIG91dHJvIHBhcmVjZSBlc3RhciBpbmNvcnJldG8uIiwiZmVlZGJhY2tfdHlwZV9wdCI6InRleHQiLCJmZWVkYmFja196aCI6IuWFtuS4reS4gOS4quWAvOaYr+ato+ehrueahO+8jOS9huWPpuS4gOS4quWAvOS8vOS5juS4jeato+ehruOAgiIsImZlZWRiYWNrX3R5cGVfemgiOiJ0ZXh0IiwiZmVlZGJhY2tfa28iOiLtlZwg6rCS7J2AIOunnuyngOunjCDri6Trpbgg6rCS7J2AIO2LgOumsCDqsoMg6rCZ7Iq164uI64ukLiIsImZlZWRiYWNrX3R5cGVfa28iOiJ0ZXh0IiwiZmVlZGJhY2tfcnUiOiLQntC00L3QviDQuNC3INC30L3QsNGH0LXQvdC40Lkg0LLQtdGA0L3Qviwg0L3QviDQtNGA0YPQs9C+0LUsINC/0L4t0LLQuNC00LjQvNC+0LzRgywg0L3QtdCy0LXRgNC90L4uIiwiZmVlZGJhY2tfdHlwZV9ydSI6InRleHQiLCJmZWVkYmFja19zdiI6IkRldCBlbmEgdsOkcmRldCDDpHIga29ycmVrdCwgbWVuIGRldCBhbmRyYSB2ZXJrYXIgdmFyYSBmZWxha3RpZ3QuIiwiZmVlZGJhY2tfdHlwZV9zdiI6InRleHQiLCJmZWVkYmFja19lcyI6IlVubyBkZSBsb3MgdmFsb3JlcyBlcyBjb3JyZWN0bywgcGVybyBlbCBvdHJvIHBhcmVjZSBzZXIgaW5jb3JyZWN0by4iLCJmZWVkYmFja190eXBlX2VzIjoidGV4dCJ9LHsicGF0dGVybiI6IjAyIiwidHJ1dGgiOiJXIiwiY2hvaWNlX2VuIjoiTGlzdEJMMiIsImNob2ljZV90eXBlX2VuIjoiY2FzIiwiY2hvaWNlX2xpc3RfZXhwcl9lbiI6dHJ1ZSwiY2hvaWNlX2xhbmd1YWdlX2luZGVwZW5kZW50Ijp0cnVlLCJjaG9pY2VfZnIiOiJMaXN0QkwyIiwiY2hvaWNlX3R5cGVfZnIiOiJjYXMiLCJjaG9pY2VfbGlzdF9leHByX2ZyIjp0cnVlLCJjaG9pY2VfZGUiOiJMaXN0QkwyIiwiY2hvaWNlX3R5cGVfZGUiOiJjYXMiLCJjaG9pY2VfbGlzdF9leHByX2RlIjp0cnVlLCJjaG9pY2VfaXQiOiJMaXN0QkwyIiwiY2hvaWNlX3R5cGVfaXQiOiJjYXMiLCJjaG9pY2VfbGlzdF9leHByX2l0Ijp0cnVlLCJjaG9pY2VfamEiOiJMaXN0QkwyIiwiY2hvaWNlX3R5cGVfamEiOiJjYXMiLCJjaG9pY2VfbGlzdF9leHByX2phIjp0cnVlLCJmZWVkYmFja19lbiI6Ik9uZSBvZiB0aGUgdmFsdWVzIGlzIGNvcnJlY3QsIGJ1dCB0aGUgb3RoZXIgYXBwZWFycyB0byBiZSBpbmNvcnJlY3QuIiwiZmVlZGJhY2tfdHlwZV9lbiI6InRleHQiLCJmZWVkYmFja19mciI6IkwndW5lIGRlcyB2YWxldXJzIGVzdCBjb3JyZWN0ZSwgbWFpcyBsJ2F1dHJlIHNlbWJsZSBpbmNvcnJlY3RlLiIsImZlZWRiYWNrX3R5cGVfZnIiOiJ0ZXh0IiwiZmVlZGJhY2tfZGUiOiJFaW5lciBkZXIgV2VydGUgaXN0IHJpY2h0aWcsIGFiZXIgZGVyIGFuZGVyZSBzY2hlaW50IGZhbHNjaCB6dSBzZWluLiIsImZlZWRiYWNrX3R5cGVfZGUiOiJ0ZXh0IiwiZmVlZGJhY2tfaXQiOiJVbm8gZGVpIHZhbG9yaSDDqCBjb3JyZXR0bywgbWEgbCdhbHRybyBzZW1icmEgZXNzZXJlIGVycmF0by4iLCJmZWVkYmFja190eXBlX2l0IjoidGV4dCIsImZlZWRiYWNrX2phIjoi5LiA5pa544Gu5YCk44Gv5q2j44GX44GE44Gu44Gn44GZ44GMLCDjgoLjgYbkuIDmlrnjgYzpgZXjgYbjgojjgYbjgafjgZkuIiwiZmVlZGJhY2tfdHlwZV9qYSI6InRleHQiLCJmZWVkYmFja19ieV90cnV0aCI6dHJ1ZSwiY2hvaWNlX3B0IjoiTGlzdEJMMiIsImNob2ljZV90eXBlX3B0IjoiY2FzIiwiY2hvaWNlX2xpc3RfZXhwcl9wdCI6dHJ1ZSwiY2hvaWNlX3poIjoiTGlzdEJMMiIsImNob2ljZV90eXBlX3poIjoiY2FzIiwiY2hvaWNlX2xpc3RfZXhwcl96aCI6dHJ1ZSwiY2hvaWNlX2tvIjoiTGlzdEJMMiIsImNob2ljZV90eXBlX2tvIjoiY2FzIiwiY2hvaWNlX2xpc3RfZXhwcl9rbyI6dHJ1ZSwiY2hvaWNlX3J1IjoiTGlzdEJMMiIsImNob2ljZV90eXBlX3J1IjoiY2FzIiwiY2hvaWNlX2xpc3RfZXhwcl9ydSI6dHJ1ZSwiY2hvaWNlX3N2IjoiTGlzdEJMMiIsImNob2ljZV90eXBlX3N2IjoiY2FzIiwiY2hvaWNlX2xpc3RfZXhwcl9zdiI6dHJ1ZSwiY2hvaWNlX2VzIjoiTGlzdEJMMiIsImNob2ljZV90eXBlX2VzIjoiY2FzIiwiY2hvaWNlX2xpc3RfZXhwcl9lcyI6dHJ1ZSwiZmVlZGJhY2tfcHQiOiJVbSBkb3MgdmFsb3JlcyBlc3TDoSBjb3JyZXRvLCBtYXMgbyBvdXRybyBwYXJlY2UgZXN0YXIgaW5jb3JyZXRvLiIsImZlZWRiYWNrX3R5cGVfcHQiOiJ0ZXh0IiwiZmVlZGJhY2tfemgiOiLlhbbkuK3kuIDkuKrlgLzmmK/mraPnoa7nmoTvvIzkvYblj6bkuIDkuKrlgLzkvLzkuY7kuI3mraPnoa7jgIIiLCJmZWVkYmFja190eXBlX3poIjoidGV4dCIsImZlZWRiYWNrX2tvIjoi7ZWcIOqwkuydgCDrp57sp4Drp4wg64uk66W4IOqwkuydgCDti4DrprAg6rKDIOqwmeyKteuLiOuLpC4iLCJmZWVkYmFja190eXBlX2tvIjoidGV4dCIsImZlZWRiYWNrX3J1Ijoi0J7QtNC90L4g0LjQtyDQt9C90LDRh9C10L3QuNC5INCy0LXRgNC90L4sINC90L4g0LTRgNGD0LPQvtC1LCDQv9C+LdCy0LjQtNC40LzQvtC80YMsINC90LXQstC10YDQvdC+LiIsImZlZWRiYWNrX3R5cGVfcnUiOiJ0ZXh0IiwiZmVlZGJhY2tfc3YiOiJEZXQgZW5hIHbDpHJkZXQgw6RyIGtvcnJla3QsIG1lbiBkZXQgYW5kcmEgdmVya2FyIHZhcmEgZmVsYWt0aWd0LiIsImZlZWRiYWNrX3R5cGVfc3YiOiJ0ZXh0IiwiZmVlZGJhY2tfZXMiOiJVbm8gZGUgbG9zIHZhbG9yZXMgZXMgY29ycmVjdG8sIHBlcm8gZWwgb3RybyBwYXJlY2Ugc2VyIGluY29ycmVjdG8uIiwiZmVlZGJhY2tfdHlwZV9lcyI6InRleHQifSx7InBhdHRlcm4iOiIwMyIsInRydXRoIjoiVyIsImNob2ljZV9lbiI6Ikxpc3RCTDMiLCJjaG9pY2VfdHlwZV9lbiI6ImNhcyIsImNob2ljZV9saXN0X2V4cHJfZW4iOnRydWUsImNob2ljZV9sYW5ndWFnZV9pbmRlcGVuZGVudCI6dHJ1ZSwiY2hvaWNlX2ZyIjoiTGlzdEJMMyIsImNob2ljZV90eXBlX2ZyIjoiY2FzIiwiY2hvaWNlX2xpc3RfZXhwcl9mciI6dHJ1ZSwiY2hvaWNlX2RlIjoiTGlzdEJMMyIsImNob2ljZV90eXBlX2RlIjoiY2FzIiwiY2hvaWNlX2xpc3RfZXhwcl9kZSI6dHJ1ZSwiY2hvaWNlX2l0IjoiTGlzdEJMMyIsImNob2ljZV90eXBlX2l0IjoiY2FzIiwiY2hvaWNlX2xpc3RfZXhwcl9pdCI6dHJ1ZSwiY2hvaWNlX2phIjoiTGlzdEJMMyIsImNob2ljZV90eXBlX2phIjoiY2FzIiwiY2hvaWNlX2xpc3RfZXhwcl9qYSI6dHJ1ZSwiZmVlZGJhY2tfZW4iOiJCb3RoIHZhbHVlcyBhcHBlYXIgdG8gYmUgaW5jb3JyZWN0LiIsImZlZWRiYWNrX3R5cGVfZW4iOiJ0ZXh0IiwiZmVlZGJhY2tfZnIiOiJMZXMgZGV1eCB2YWxldXJzIHNlbWJsZW50IGluY29ycmVjdGVzLiIsImZlZWRiYWNrX3R5cGVfZnIiOiJ0ZXh0IiwiZmVlZGJhY2tfZGUiOiJCZWlkZSBXZXJ0ZSBzY2hlaW5lbiBmYWxzY2ggenUgc2Vpbi4iLCJmZWVkYmFja190eXBlX2RlIjoidGV4dCIsImZlZWRiYWNrX2l0IjoiRW50cmFtYmkgaSB2YWxvcmkgc2VtYnJhbm8gZXNzZXJlIGVycmF0aS4iLCJmZWVkYmFja190eXBlX2l0IjoidGV4dCIsImZlZWRiYWNrX2phIjoi44Gp44Gh44KJ44Gu5YCk44KC6YGV44GG44KI44GG44Gn44GZLiIsImZlZWRiYWNrX3R5cGVfamEiOiJ0ZXh0IiwiZmVlZGJhY2tfYnlfdHJ1dGgiOnRydWUsImNob2ljZV9wdCI6Ikxpc3RCTDMiLCJjaG9pY2VfdHlwZV9wdCI6ImNhcyIsImNob2ljZV9saXN0X2V4cHJfcHQiOnRydWUsImNob2ljZV96aCI6Ikxpc3RCTDMiLCJjaG9pY2VfdHlwZV96aCI6ImNhcyIsImNob2ljZV9saXN0X2V4cHJfemgiOnRydWUsImNob2ljZV9rbyI6Ikxpc3RCTDMiLCJjaG9pY2VfdHlwZV9rbyI6ImNhcyIsImNob2ljZV9saXN0X2V4cHJfa28iOnRydWUsImNob2ljZV9ydSI6Ikxpc3RCTDMiLCJjaG9pY2VfdHlwZV9ydSI6ImNhcyIsImNob2ljZV9saXN0X2V4cHJfcnUiOnRydWUsImNob2ljZV9zdiI6Ikxpc3RCTDMiLCJjaG9pY2VfdHlwZV9zdiI6ImNhcyIsImNob2ljZV9saXN0X2V4cHJfc3YiOnRydWUsImNob2ljZV9lcyI6Ikxpc3RCTDMiLCJjaG9pY2VfdHlwZV9lcyI6ImNhcyIsImNob2ljZV9saXN0X2V4cHJfZXMiOnRydWUsImZlZWRiYWNrX3B0IjoiQW1ib3Mgb3MgdmFsb3JlcyBwYXJlY2VtIGVzdGFyIGluY29ycmV0b3MuIiwiZmVlZGJhY2tfdHlwZV9wdCI6InRleHQiLCJmZWVkYmFja196aCI6IuS4pOS4quWAvOS8vOS5jumDveS4jeato+ehruOAgiIsImZlZWRiYWNrX3R5cGVfemgiOiJ0ZXh0IiwiZmVlZGJhY2tfa28iOiLrkZAg6rCSIOuqqOuRkCDti4DrprAg6rKDIOqwmeyKteuLiOuLpC4iLCJmZWVkYmFja190eXBlX2tvIjoidGV4dCIsImZlZWRiYWNrX3J1Ijoi0J7QsdCwINC30L3QsNGH0LXQvdC40Y8sINC/0L4t0LLQuNC00LjQvNC+0LzRgywg0L3QtdCy0LXRgNC90YsuIiwiZmVlZGJhY2tfdHlwZV9ydSI6InRleHQiLCJmZWVkYmFja19zdiI6IkLDpWRhIHbDpHJkZW5hIHZlcmthciB2YXJhIGZlbGFrdGlnYS4iLCJmZWVkYmFja190eXBlX3N2IjoidGV4dCIsImZlZWRiYWNrX2VzIjoiQW1ib3MgdmFsb3JlcyBwYXJlY2VuIHNlciBpbmNvcnJlY3Rvcy4iLCJmZWVkYmFja190eXBlX2VzIjoidGV4dCJ9XSwicXZhcnMiOiIvKioqKioqKioqKioqKiogR2F1c3NFbGltUm93UmVkdWNlZENoZWNrLUEgKHJiKSAqKioqKioqKioqKioqKiovXG4vKiAwMDEuTUNCIGluY2x1ZGUgZmlsZVx0XHQgICAgICAgICAgICAgICAgICAgICAgICAgICAgICAqL1xuLyogT2JqZWN0aXZlOiBUbyBsZWFybi9jaGVjayBob3cgdG8gY2hlY2sgcm93IHJlZHVjZWQgZWNoZWxvbiAqL1xuLyogICAgICAgICAgICByYW5rOjMgICAgICAgICAgICAgICAgICAgICAgICAgICAgICAgICAgICAgICAgICAqL1xuLyoqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqL1xuLyogYzpbMV0sIHc6WzQsIDQsIDRdOyAqL1xuaWYgbm90IG51bWJlcnAoJV9NQ1FfTlVNX09QVFMpIG9yICVfTUNRX05VTV9PUFRTPjEzIHRoZW4gJV9NQ1FfTlVNX09QVFM6NztcbiVfTUNRX05VTV9DT1BUUzoxO1xuLyoqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqL1xuJV9ycChMKTo9cmFuZG9tX3Blcm11dGF0aW9uKEwpO1xudGV4MkwoTCk6PWJsb2NrKFtzOlwiXFxcXChcIiwga10sIGZvciBrOjEgdGhydSBsZW5ndGgoTCkgZG8gczpzY29uY2F0KHMsIGlmIHN0cmluZ3AoTFtrXSkgdGhlbiBMW2tdIGVsc2UgdGV4MShMW2tdKSwgXCIgXCIpLCBzOnNjb25jYXQocywgXCJcXFxcKVwiKSwgcmV0dXJuKHMpKTtcbmFMOm1ha2VfcmFuZG9tX2Rpc3RyaWJ1dGlvbl9wbShbWzEsIDRdLCBbMiwgM10sIFszLCAyXSwgWzQsIDFdLCBbNSwgMV1dKTtcbiVfbWttYXQociwgYyk6PWdlbm1hdHJpeChsYW1iZGEoW2ksIGpdLCByYW5kKGFMKSksIHIsIGMpO1xuJV9ubWtfYWxsOiVfcnAobGlzdGlmeShjYXJ0ZXNpYW5fcHJvZHVjdCh7MSwgMiwgMywgNCwgNX0sIHsxLCAyLCAzLCA0LCA1fSkpKTtcbm5tTDpbMSwgMiwgMywgNCwgNV07XG4lX19sYW5nKEwpOj0lX19tY3FfbGFuZyhMLCAlX1NUQUNLX0xBTkcpO1xuXG5pZiBub3QgbnVtYmVycCglX25yKSB0aGVuICVfbnI6NDtcbmlmIG5vdCBudW1iZXJwKCVfbmMpIHRoZW4gJV9uYzo1O1xuaWYgbm90IGxpc3RwKCVfcHYpIHRoZW4gJV9wdjpyYW5kKFtbMSwgMiwgM10sIFsxLCAzLCA0XSwgWzEsIDIsIDRdXSk7XG5cbiVfcHZDOmxpc3RpZnkoc2V0ZGlmZmVyZW5jZShzZXRpZnkobWFrZWxpc3QoaywgaywgJV9uYykpLCBzZXRpZnkoJV9wdikpKTtcblxubWF0RjpyYW5kZWNoZWxvbiglX25yLCAlX25jLCAlX3B2LCBbLTMsIC0yLCAtMSwgMSwgMiwgM10pO1xucEw6bWFrZV9yYW5kb21fZGlzdHJpYnV0aW9uX3BtKFtbMywgMV0sIFsyLCAzXSwgWzEsIDEwXSwgWzAsIDNdXSk7XG5tYXRQOnJhbmRvbV9tYXRyaXhfZGV0X2xpc3QoJV9uciwgcEwsIFstMSwgMV0pO1xubWF0QTptYXRQLm1hdEY7XG5wcDptYXRBWzIsICVfcHZDWzJdXTtcbnFxOm1hdEZbMSwgJV9wdkNbMV1dO1xuXG5tYXRGcTpnZW5tYXRyaXgobGFtYmRhKFtpLCBqXSwgaWYgW2ksIGpdPVsxLCAlX3B2Q1sxXV0gdGhlbiBxIGVsc2UgbWF0RltpLCBqXSksICVfbnIsICVfbmMpO1xubWF0QXA6Z2VubWF0cml4KGxhbWJkYShbaSwgal0sIGlmIFtpLCBqXT1bMiwgJV9wdkNbMl1dIHRoZW4gcCBlbHNlIG1hdEFbaSwgal0pLCAlX25yLCAlX25jKTtcblxuJV9ta29wKHAsIHEpOj1jYXN0ZXh0KFwiXFwocD17QHBAfSwgcT17QHFAfVxcKVwiKTtcbkxpc3RBTDE6WyVfbWtvcChwcCwgcXEpXTtcblxuLyogd3JvbmcgbXVsdGlwbGljYXRpb24gKi9cbkxpc3RCTDE6WyVfbWtvcChwcCsxLCBxcSldO1xuTGlzdEJMMTplbmRjb25zKCVfbWtvcChwcC0xLCBxcSksIExpc3RCTDEpO1xuTGlzdEJMMTplbmRjb25zKCVfbWtvcChwcC0yLCBxcSksIExpc3RCTDEpO1xuTGlzdEJMMTplbmRjb25zKCVfbWtvcChwcCsyLCBxcSksIExpc3RCTDEpO1xuXG5MaXN0QkwyOlslX21rb3AocHAsIHFxKzEpXTtcbkxpc3RCTDI6ZW5kY29ucyglX21rb3AocHAsIHFxLTEpLCBMaXN0QkwyKTtcbkxpc3RCTDI6ZW5kY29ucyglX21rb3AocHAsIHFxKzIpLCBMaXN0QkwyKTtcbkxpc3RCTDI6ZW5kY29ucyglX21rb3AocHAsIHFxLTIpLCBMaXN0QkwyKTtcblxuTGlzdEJMMzpbJV9ta29wKHBwKzEsIHFxLTEpXTtcbkxpc3RCTDM6ZW5kY29ucyglX21rb3AocHAtMSwgcXErMSksIExpc3RCTDMpO1xuTGlzdEJMMzplbmRjb25zKCVfbWtvcChwcC0yLCBxcStyYW5kKFsxLCAyXSkpLCBMaXN0QkwzKTtcbkxpc3RCTDM6ZW5kY29ucyglX21rb3AocHArMiwgcXEtcmFuZChbMSwgMl0pKSwgTGlzdEJMMyk7XG5cbi8qKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqIFF1ZXN0aW9uIHRleHQgKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKi9cbi8qICVfX21jcV9xdGV4dEwgOiBxdWVzdGlvbiB0ZXh0IGxhbmd1YWdlIGxpc3QgICAgICAgICAgICAgICAgICAgICAgICAgICAgICAgICAgICAgICAgICAgKi9cbi8qIE5PVEU6IEluIHRoZSBxdWVzdGlvbiB0ZXh0LCBzdHJpbmcgXCJfX1NFTFRZUEVfX1wiIGlzIHJlcGxhY2VkIGF1dG9tYXRpY2FsbHkuICAgICAgICAgICAqL1xuLyoqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqL1xuLyogaWYgbm90IGRlcGVuZCBvbiBsb2NhbGUsIGRlZmluZSAnJV9fbWNxX3F0ZXh0JyhtYXkgaW5jbHVkZSBfX1NFTFRZUEVfXykgb3IgJ3F0ZXh0JyBkaXJlY3RseSAqL1xuLyogTUNRX1dFQkFQUF9TTE9UOnF0ZXh0ICovXG5cbi8qKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKiBMaXN0IG9mIHBhdHRlcm5zL2ZlZWRiYWNrcyAqKioqKioqKioqKioqKioqKioqKioqKioqKioqKi9cbi8qICBUaGlzIHRlbXBsYXRlIGFzc3VtZXMgdGhlIGxpc3RzICVfX0NvcHQ/TCwgJV9fV29wdD9MIGRvZXMgbm90IGRlcGVuZCBvbiB0aGUgbG9jYWxlLiAgKi9cbi8qICBJZiBpdCBkZXBlbmRzIG9uIHRoZSBsb2NhbGUsIHBsZWFzZSBzZXQgdXAgdGhlIGxpc3QgdXNpbmcgJV9TVEFDS19MQU5HLiAgICAgICAgICAgICAgKi9cbi8qKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKioqKi9cbi8qIENvcnJlY3QgcGF0dGVybiAxICovXG4vKiBNQ1FfV0VCQVBQX1NMT1Q6Q29wdDEgKi9cbi8qIE1DUV9XRUJBUFBfU0xPVDpDbXNnMSAqL1xuXG4vKiBXcm9uZyBhbnN3ZXIgcGF0dGVybiAxICovXG4vKiBNQ1FfV0VCQVBQX1NMT1Q6V29wdDEgKi9cbi8qIE1DUV9XRUJBUFBfU0xPVDpXbXNnMSAqL1xuXG4vKiBXcm9uZyBhbnN3ZXIgcGF0dGVybiAxICovXG4vKiBNQ1FfV0VCQVBQX1NMT1Q6V29wdDIgKi9cbi8qIE1DUV9XRUJBUFBfU0xPVDpXbXNnMiAqL1xuXG4vKiBXcm9uZyBhbnN3ZXIgcGF0dGVybiAxICovXG4vKiBNQ1FfV0VCQVBQX1NMT1Q6V29wdDMgKi9cbi8qIE1DUV9XRUJBUFBfU0xPVDpXbXNnMyAqL1xuXG5cbi8qKioqKioqKioqIEVPRiAqKioqKioqKioqL1xuIiwicGFyYW1ldGVycyI6IiVfbnI6MztcbiVfbmM6NDtcbiVfcHY6cmFuZChbWzEsIDJdLCBbMSwgM11dKTsiLCJzZXR0aW5ncyI6eyJub0NvcnJlY3RPcHRpb24iOmZhbHNlLCJub0lkZWFPcHRpb24iOmZhbHNlLCJyYWRpb011bHRpcGxlUHJvbXB0IjpmYWxzZSwiY2FzdGV4dFRlbXBsYXRlIjp0cnVlLCJzY29yaW5nTWV0aG9kIjoiMSIsIm51bU9wdGlvbnMiOiI3IiwibnVtQ29ycmVjdCI6IjEiLCJyYW5kb21Db3JyZWN0IjpmYWxzZSwiY29ycmVjdENvdW50cyI6IjEsIDIiLCJyZXF1aXJlUGFpcnMiOmZhbHNlLCJmZWVkYmFja0J5VHJ1dGgiOmZhbHNlLCJmZWVkYmFja01vZGUiOiJ0cnVlIn0sImxlZ2FjeVF1ZXN0aW9uSW5wdXRzIjp7ImVuIjp7Im9yaWdpbmFsIjoic2NvbmNhdChcIlRoZSBcIiwgdGV4MkwoWyVfbnIsIFwiXFxcXHRpbWVzXCIsICVfbmNdKSwgXCIgbWF0cml4IFxcXFwoQT1cXFxcKVwiLCB0ZXgyKG1hdEEpLCBcIiBpcyB0cmFuc2Zvcm1lZCBpbnRvIFwiLCB0ZXgyKG1hdEZRKSwgXCIgYnkgZWxlbWVudGFyeSByb3cgb3BlcmF0aW9ucy4gQ2hvb3NlIHRoZSBjb3JyZWN0IHBhaXIgb2YgdmFsdWVzIG9mIFxcXFwocCwgcVxcXFwpLlwiKSIsImNvbnZlcnRlZCI6IlRoZSBcXCh7QCVfbnJAfSBcXHRpbWVzIHtAJV9uY0B9XFwpIG1hdHJpeCBcXChBPVxcKXtAbWF0QUB9IGlzIHRyYW5zZm9ybWVkIGludG8ge0BtYXRGUUB9IGJ5IGVsZW1lbnRhcnkgcm93IG9wZXJhdGlvbnMuIENob29zZSB0aGUgY29ycmVjdCBwYWlyIG9mIHZhbHVlcyBvZiBcXChwLCBxXFwpLiJ9LCJmciI6eyJvcmlnaW5hbCI6InNjb25jYXQoXCJMYSBtYXRyaWNlIFwiLCB0ZXgyTChbJV9uciwgXCJcXFxcdGltZXNcIiwgJV9uY10pLCBcIiBcXFxcKEE9XFxcXClcIiwgdGV4MihtYXRBKSwgXCIgZXN0IHRyYW5zZm9ybcOpZSBlbiBcIiwgdGV4MihtYXRGUSksIFwiIHBhciBkZXMgb3DDqXJhdGlvbnMgw6lsw6ltZW50YWlyZXMgc3VyIGxlcyBsaWduZXMuIENob2lzaXNzZXogbGUgY291cGxlIGNvcnJlY3QgZGUgdmFsZXVycyBkZSBcXFxcKHAsIHFcXFxcKS5cIikiLCJjb252ZXJ0ZWQiOiJMYSBtYXRyaWNlIFxcKHtAJV9uckB9IFxcdGltZXMge0AlX25jQH1cXCkgXFwoQT1cXCl7QG1hdEFAfSBlc3QgdHJhbnNmb3Jtw6llIGVuIHtAbWF0RlFAfSBwYXIgZGVzIG9ww6lyYXRpb25zIMOpbMOpbWVudGFpcmVzIHN1ciBsZXMgbGlnbmVzLiBDaG9pc2lzc2V6IGxlIGNvdXBsZSBjb3JyZWN0IGRlIHZhbGV1cnMgZGUgXFwocCwgcVxcKS4ifSwiaXQiOnsib3JpZ2luYWwiOiJzY29uY2F0KFwiTGEgbWF0cmljZSBcIiwgdGV4MkwoWyVfbnIsIFwiXFxcXHRpbWVzXCIsICVfbmNdKSwgXCIgXFxcXChBPVxcXFwpXCIsIHRleDIobWF0QSksIFwiIHZpZW5lIHRyYXNmb3JtYXRhIGluIFwiLCB0ZXgyKG1hdEZRKSwgXCIgbWVkaWFudGUgb3BlcmF6aW9uaSBlbGVtZW50YXJpIGRpIHJpZ2EuIFNjZWdsaSBsYSBjb3BwaWEgY29ycmV0dGEgZGkgdmFsb3JpIGRpIFxcXFwocCwgcVxcXFwpLlwiKSIsImNvbnZlcnRlZCI6IkxhIG1hdHJpY2UgXFwoe0AlX25yQH0gXFx0aW1lcyB7QCVfbmNAfVxcKSBcXChBPVxcKXtAbWF0QUB9IHZpZW5lIHRyYXNmb3JtYXRhIGluIHtAbWF0RlFAfSBtZWRpYW50ZSBvcGVyYXppb25pIGVsZW1lbnRhcmkgZGkgcmlnYS4gU2NlZ2xpIGxhIGNvcHBpYSBjb3JyZXR0YSBkaSB2YWxvcmkgZGkgXFwocCwgcVxcKS4ifSwiZGUiOnsib3JpZ2luYWwiOiJzY29uY2F0KFwiRGllIFwiLCB0ZXgyTChbJV9uciwgXCJcXFxcdGltZXNcIiwgJV9uY10pLCBcIi1NYXRyaXggXFxcXChBPVxcXFwpXCIsIHRleDIobWF0QSksIFwiIHdpcmQgZHVyY2ggZWxlbWVudGFyZSBaZWlsZW51bWZvcm11bmdlbiBpbiBcIiwgdGV4MihtYXRGUSksIFwiIHVtZ2Vmb3JtdC4gV8OkaGxlbiBTaWUgZGFzIHJpY2h0aWdlIFdlcnRlcGFhciBmw7xyIFxcXFwocCwgcVxcXFwpLlwiKSIsImNvbnZlcnRlZCI6IkRpZSBcXCh7QCVfbnJAfSBcXHRpbWVzIHtAJV9uY0B9XFwpLU1hdHJpeCBcXChBPVxcKXtAbWF0QUB9IHdpcmQgZHVyY2ggZWxlbWVudGFyZSBaZWlsZW51bWZvcm11bmdlbiBpbiB7QG1hdEZRQH0gdW1nZWZvcm10LiBXw6RobGVuIFNpZSBkYXMgcmljaHRpZ2UgV2VydGVwYWFyIGbDvHIgXFwocCwgcVxcKS4ifX0sImluY2x1ZGVGaWxlbmFtZSI6IkdhdXNzRWxpbVJvd1JlZHVjZWRDaGVjay1BLnR4dCIsImluY2x1ZGVTb3VyY2UiOm51bGx9 */

/******************************** MAIN QUESTION VARIABLES ********************************/


%_MCQ_NUM_OPTS:7;
%_MCQ_NUM_COPTS:1;
%__mcq_nocorrectopt:false;
%__mcq_noidea:false;
%__mcq_scmethod:1;
%_nr:3;
%_nc:4;
%_pv:rand([[1, 2], [1, 3]]);

/**************** Generated by mcq-webapp ****************/
/************** GaussElimRowReducedCheck-A (rb) ***************/
/* 001.MCB include file		                              */
/* Objective: To learn/check how to check row reduced echelon */
/*            rank:3                                          */
/**************************************************************/
/* c:[1], w:[4, 4, 4]; */
if not numberp(%_MCQ_NUM_OPTS) or %_MCQ_NUM_OPTS>13 then %_MCQ_NUM_OPTS:7;
%_MCQ_NUM_COPTS:1;
/********************************************************************/
%_rp(L):=random_permutation(L);
tex2L(L):=block([s:"\\(", k], for k:1 thru length(L) do s:sconcat(s, if stringp(L[k]) then L[k] else tex1(L[k]), " "), s:sconcat(s, "\\)"), return(s));
aL:make_random_distribution_pm([[1, 4], [2, 3], [3, 2], [4, 1], [5, 1]]);
%_mkmat(r, c):=genmatrix(lambda([i, j], rand(aL)), r, c);
%_nmk_all:%_rp(listify(cartesian_product({1, 2, 3, 4, 5}, {1, 2, 3, 4, 5})));
nmL:[1, 2, 3, 4, 5];
%__lang(L):=%__mcq_lang(L, %_STACK_LANG);

if not numberp(%_nr) then %_nr:4;
if not numberp(%_nc) then %_nc:5;
if not listp(%_pv) then %_pv:rand([[1, 2, 3], [1, 3, 4], [1, 2, 4]]);

%_pvC:listify(setdifference(setify(makelist(k, k, %_nc)), setify(%_pv)));

matF:randechelon(%_nr, %_nc, %_pv, [-3, -2, -1, 1, 2, 3]);
pL:make_random_distribution_pm([[3, 1], [2, 3], [1, 10], [0, 3]]);
matP:random_matrix_det_list(%_nr, pL, [-1, 1]);
matA:matP.matF;
pp:matA[2, %_pvC[2]];
qq:matF[1, %_pvC[1]];

matFq:genmatrix(lambda([i, j], if [i, j]=[1, %_pvC[1]] then q else matF[i, j]), %_nr, %_nc);
matAp:genmatrix(lambda([i, j], if [i, j]=[2, %_pvC[2]] then p else matA[i, j]), %_nr, %_nc);

%_mkop(p, q):=castext("\(p={@p@}, q={@q@}\)");
ListAL1:[%_mkop(pp, qq)];

/* wrong multiplication */
ListBL1:[%_mkop(pp+1, qq)];
ListBL1:endcons(%_mkop(pp-1, qq), ListBL1);
ListBL1:endcons(%_mkop(pp-2, qq), ListBL1);
ListBL1:endcons(%_mkop(pp+2, qq), ListBL1);

ListBL2:[%_mkop(pp, qq+1)];
ListBL2:endcons(%_mkop(pp, qq-1), ListBL2);
ListBL2:endcons(%_mkop(pp, qq+2), ListBL2);
ListBL2:endcons(%_mkop(pp, qq-2), ListBL2);

ListBL3:[%_mkop(pp+1, qq-1)];
ListBL3:endcons(%_mkop(pp-1, qq+1), ListBL3);
ListBL3:endcons(%_mkop(pp-2, qq+rand([1, 2])), ListBL3);
ListBL3:endcons(%_mkop(pp+2, qq-rand([1, 2])), ListBL3);

/********************************** Question text ****************************************/
/* %__mcq_qtextL : question text language list                                           */
/* NOTE: In the question text, string "__SELTYPE__" is replaced automatically.           */
/*****************************************************************************************/
/* if not depend on locale, define '%__mcq_qtext'(may include __SELTYPE__) or 'qtext' directly */
/* MCQ_WEBAPP_SLOT:qtext */

/******************************** List of patterns/feedbacks *****************************/
/*  This template assumes the lists %__Copt?L, %__Wopt?L does not depend on the locale.  */
/*  If it depends on the locale, please set up the list using %_STACK_LANG.              */
/*****************************************************************************************/
/* Correct pattern 1 */
/* MCQ_WEBAPP_SLOT:Copt1 */
/* MCQ_WEBAPP_SLOT:Cmsg1 */

/* Wrong answer pattern 1 */
/* MCQ_WEBAPP_SLOT:Wopt1 */
/* MCQ_WEBAPP_SLOT:Wmsg1 */

/* Wrong answer pattern 1 */
/* MCQ_WEBAPP_SLOT:Wopt2 */
/* MCQ_WEBAPP_SLOT:Wmsg2 */

/* Wrong answer pattern 1 */
/* MCQ_WEBAPP_SLOT:Wopt3 */
/* MCQ_WEBAPP_SLOT:Wmsg3 */


/********** EOF **********/

%__mcq_qtextL:[["en", castext("Applying elementary row operations to the {@%_nr@}\\times{@%_nc@} matrix \\(A=\\) {@matAp@} yields {@matFq@}. Select the correct pair of values for \\(p, q\\).")],
 ["ja", castext("{@%_nr@}\\times{@%_nc@} 行列\\(A=\\) {@matAp@} を行に関して基本変形すると, {@matFq@} になる. このとき, \\(p, q\\)の値の組として正しいものを選べ．")],
 ["fr", castext("L'application d'opérations élémentaires sur les lignes de la matrice {@%_nr@}\\times{@%_nc@} \\(A=\\) {@matAp@} donne {@matFq@}. Sélectionnez le couple correct de valeurs de \\(p, q\\).")],
 ["it", castext("Applicando operazioni elementari sulle righe alla matrice {@%_nr@}\\times{@%_nc@} \\(A=\\) {@matAp@}, si ottiene {@matFq@}. Seleziona la coppia corretta di valori di \\(p, q\\).")],
 ["de", castext("Durch elementare Zeilenumformungen wird die {@%_nr@}\\times{@%_nc@}-Matrix \\(A=\\) {@matAp@} in {@matFq@} überführt. Wählen Sie das richtige Wertepaar für \\(p, q\\) aus.")],
 ["pt", castext("Aplicando operações elementares sobre as linhas à matriz {@%_nr@}\\times{@%_nc@} \\(A=\\) {@matAp@}, obtém-se {@matFq@}. Selecione o par correto de valores de \\(p, q\\).")],
 ["zh", castext("对{@%_nr@}\\times{@%_nc@}矩阵\\(A=\\) {@matAp@}进行初等行变换后，得到{@matFq@}。请选择\\(p, q\\)的正确取值组合。")],
 ["ko", castext("{@%_nr@}\\times{@%_nc@} 행렬 \\(A=\\) {@matAp@}에 기본 행 연산을 적용하면 {@matFq@}가 된다. 이때 \\(p, q\\)의 값으로 올바른 순서쌍을 고르시오.")],
 ["ru", castext("После элементарных преобразований строк матрица размера {@%_nr@}\\times{@%_nc@} \\(A=\\) {@matAp@} принимает вид {@matFq@}. Выберите правильную пару значений \\(p, q\\).")],
 ["sv", castext("När elementära radoperationer utförs på {@%_nr@}\\times{@%_nc@}-matrisen \\(A=\\) {@matAp@}, erhålls {@matFq@}. Välj det korrekta värdeparet för \\(p, q\\).")],
 ["es", castext("Al aplicar operaciones elementales por filas a la matriz {@%_nr@}\\times{@%_nc@} \\(A=\\) {@matAp@}, se obtiene {@matFq@}. Selecciona el par correcto de valores de \\(p, q\\).")]];

/* Fixed true/false patterns (pair requirement disabled). */

%__CoptL1:flatten([ListAL1]);
%__Cmsg1L:[["en", castext("Correct. You seem to understand the properties of reduced row echelon matrices.")],
 ["ja", castext("正解です. 被約階段行列(行簡約行列)の性質を理解しているようです.")],
 ["fr", castext("Correct. Vous semblez comprendre les propriétés des matrices sous forme échelonnée réduite.")],
 ["it", castext("Corretto. Sembra che tu abbia compreso le proprietà delle matrici in forma a scala ridotta per righe.")],
 ["de", castext("Richtig. Sie scheinen die Eigenschaften von Matrizen in reduzierter Zeilenstufenform verstanden zu haben.")],
 ["pt", castext("Correto. Parece que compreende as propriedades das matrizes na forma escalonada reduzida por linhas.")],
 ["zh", castext("回答正确。看来你已经理解了简化行阶梯形矩阵的性质。")],
 ["ko", castext("정답입니다. 기약 행 사다리꼴 행렬의 성질을 잘 이해하고 있는 것 같습니다.")],
 ["ru", castext("Верно. По-видимому, вы понимаете свойства матриц в приведённой ступенчатой форме.")],
 ["sv", castext("Rätt. Du verkar förstå egenskaperna hos matriser på reducerad trappstegsform.")],
 ["es", castext("Correcto. Parece que comprendes las propiedades de las matrices en forma escalonada reducida por filas.")]];

%__WoptL1:flatten([ListBL1]);
%__Wmsg1L:[["en", castext("One of the values is correct, but the other appears to be incorrect.")],
 ["ja", castext("一方の値は正しいのですが, もう一方が違うようです.")],
 ["fr", castext("L'une des valeurs est correcte, mais l'autre semble incorrecte.")],
 ["it", castext("Uno dei valori è corretto, ma l'altro sembra essere errato.")],
 ["de", castext("Einer der Werte ist richtig, aber der andere scheint falsch zu sein.")],
 ["pt", castext("Um dos valores está correto, mas o outro parece estar incorreto.")],
 ["zh", castext("其中一个值是正确的，但另一个值似乎不正确。")],
 ["ko", castext("한 값은 맞지만 다른 값은 틀린 것 같습니다.")],
 ["ru", castext("Одно из значений верно, но другое, по-видимому, неверно.")],
 ["sv", castext("Det ena värdet är korrekt, men det andra verkar vara felaktigt.")],
 ["es", castext("Uno de los valores es correcto, pero el otro parece ser incorrecto.")]];

%__WoptL2:flatten([ListBL2]);
%__Wmsg2L:[["en", castext("One of the values is correct, but the other appears to be incorrect.")],
 ["ja", castext("一方の値は正しいのですが, もう一方が違うようです.")],
 ["fr", castext("L'une des valeurs est correcte, mais l'autre semble incorrecte.")],
 ["it", castext("Uno dei valori è corretto, ma l'altro sembra essere errato.")],
 ["de", castext("Einer der Werte ist richtig, aber der andere scheint falsch zu sein.")],
 ["pt", castext("Um dos valores está correto, mas o outro parece estar incorreto.")],
 ["zh", castext("其中一个值是正确的，但另一个值似乎不正确。")],
 ["ko", castext("한 값은 맞지만 다른 값은 틀린 것 같습니다.")],
 ["ru", castext("Одно из значений верно, но другое, по-видимому, неверно.")],
 ["sv", castext("Det ena värdet är korrekt, men det andra verkar vara felaktigt.")],
 ["es", castext("Uno de los valores es correcto, pero el otro parece ser incorrecto.")]];

%__WoptL3:flatten([ListBL3]);
%__Wmsg3L:[["en", castext("Both values appear to be incorrect.")],
 ["ja", castext("どちらの値も違うようです.")],
 ["fr", castext("Les deux valeurs semblent incorrectes.")],
 ["it", castext("Entrambi i valori sembrano essere errati.")],
 ["de", castext("Beide Werte scheinen falsch zu sein.")],
 ["pt", castext("Ambos os valores parecem estar incorretos.")],
 ["zh", castext("两个值似乎都不正确。")],
 ["ko", castext("두 값 모두 틀린 것 같습니다.")],
 ["ru", castext("Оба значения, по-видимому, неверны.")],
 ["sv", castext("Båda värdena verkar vara felaktiga.")],
 ["es", castext("Ambos valores parecen ser incorrectos.")]];

/**************** End of generated variables ****************/

/***************************** END OF MAIN QUESTION VARIABLES ****************************/

/*********************************** DO NOT EDIT HERE ************************************/
/*************** patterns is automatically generate by the post template *****************/
stack_include("https://yositomi-opu.github.io/stack_questions/mcq_template_post_cas.mac");
/*****************************************************************************************/

%__mcq_langcode:castext("[[lang code=\"en\"]][[/lang]][[lang code=\"ja\"]][[/lang]][[lang code=\"fr\"]][[/lang]][[lang code=\"it\"]][[/lang]][[lang code=\"de\"]][[/lang]][[lang code=\"pt\"]][[/lang]][[lang code=\"zh\"]][[/lang]][[lang code=\"ko\"]][[/lang]][[lang code=\"ru\"]][[/lang]][[lang code=\"sv\"]][[/lang]][[lang code=\"es\"]][[/lang]]");
]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="html">
      <text>001.GaussElimRowReducedCheck-B-3x4rk2-rb</text>
    </questionnote>
    <questiondescription format="html">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span>
<span lang="ja" class="multilang">よくできました。正解です！</span>
<span lang="en" class="multilang">Well done! Spot on!</span>
<span lang="es" class="multilang">¡Muy bien! ¡Respuesta correcta!</span>
<span lang="fr" class="multilang">Bravo ! Bonne réponse !</span>
<span lang="it" class="multilang">Ben fatto! Risposta corretta!</span>
<span lang="de" class="multilang">Gut gemacht! Richtig!</span>
<span lang="pt" class="multilang">Muito bem! Resposta correta!</span>
<span lang="zh" class="multilang">做得很好！回答正确！</span>
<span lang="ko" class="multilang">잘했습니다! 정답입니다!</span>
<span lang="ru" class="multilang">Отлично! Верно!</span>
<span lang="sv" class="multilang">Bra gjort! Rätt svar!</span>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span>
<span lang="ja" class="multilang">部分的に正解です。</span>
<span lang="en" class="multilang">Partially correct!</span>
<span lang="es" class="multilang">Parcialmente correcto.</span>
<span lang="fr" class="multilang">Partiellement correct.</span>
<span lang="it" class="multilang">Parzialmente corretto.</span>
<span lang="de" class="multilang">Teilweise richtig.</span>
<span lang="pt" class="multilang">Parcialmente correto.</span>
<span lang="zh" class="multilang">部分正确。</span>
<span lang="ko" class="multilang">부분적으로 맞았습니다.</span>
<span lang="ru" class="multilang">Частично верно.</span>
<span lang="sv" class="multilang">Delvis rätt.</span>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span>
<span lang="ja" class="multilang">間違いです。</span>
<span lang="en" class="multilang">Incorrect!</span>
<span lang="es" class="multilang">Respuesta incorrecta.</span>
<span lang="fr" class="multilang">Incorrect.</span>
<span lang="it" class="multilang">Errato.</span>
<span lang="de" class="multilang">Falsch.</span>
<span lang="pt" class="multilang">Incorreto.</span>
<span lang="zh" class="multilang">回答错误。</span>
<span lang="ko" class="multilang">오답입니다.</span>
<span lang="ru" class="multilang">Неверно.</span>
<span lang="sv" class="multilang">Felaktigt.</span>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>(</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>radio</type>
      <tans>ta1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered,hideanswer</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[stack_include("https://yositomi-opu.github.io/stack_questions/mcq_template_fvar_cas.mac");
]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>NOIDEA-option-detection</description>
        <answertest>AlgEquiv</answertest>
        <sans>flgNOID</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[{@%__mcq_langcode@}
<p>{@%__mcq_noidea_checked@}</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Score and all messages are given here</description>
        <answertest>AlgEquiv</answertest>
        <sans>chk1</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<!-- ラジオボタンで正解オプションが 2個以上ある場合に対応 -->
{@%__mcq_langcode@}
<p>{@msg@}</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>sc</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<!-- general error / 一般の誤り -->
{@%__mcq_langcode@}
<p>{@msg@}. </p>
<p>Score: {@sc@}</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>10</name>
        <description>Wrong(Incorrect) Pattern 4</description>
        <answertest>AlgEquiv</answertest>
        <sans>flgWL[4]</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>11</truenextnode>
        <trueanswernote>prt1-11-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>11</falsenextnode>
        <falseanswernote>prt1-11-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>11</name>
        <description>Wrong(Incorrect) Pattern 5</description>
        <answertest>AlgEquiv</answertest>
        <sans>flgWL[5]</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>12</truenextnode>
        <trueanswernote>prt1-12-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>12</falsenextnode>
        <falseanswernote>prt1-12-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>12</name>
        <description>Wrong(Incorrect) Pattern 6</description>
        <answertest>AlgEquiv</answertest>
        <sans>flgWL[6]</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>13</truenextnode>
        <trueanswernote>prt1-13-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>13</falsenextnode>
        <falseanswernote>prt1-13-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>13</name>
        <description>Wrong(Incorrect) Pattern 7</description>
        <answertest>AlgEquiv</answertest>
        <sans>flgWL[7]</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>14</truenextnode>
        <trueanswernote>prt1-14-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>14</falsenextnode>
        <falseanswernote>prt1-14-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>14</name>
        <description>Wrong(Incorrect) Pattern 8</description>
        <answertest>AlgEquiv</answertest>
        <sans>flgWL[8]</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>15</truenextnode>
        <trueanswernote>prt1-15-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>15</falsenextnode>
        <falseanswernote>prt1-15-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>15</name>
        <description>Wrong(Incorrect) Pattern 9</description>
        <answertest>AlgEquiv</answertest>
        <sans>flgWL[9]</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-16-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-16-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Correct Pattern 1</description>
        <answertest>AlgEquiv</answertest>
        <sans>flgCL[1]</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Correct Pattern 2</description>
        <answertest>AlgEquiv</answertest>
        <sans>flgCL[2]</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>4</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>Correct Pattern 3</description>
        <answertest>AlgEquiv</answertest>
        <sans>flgCL[3]</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>5</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>5</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>5</name>
        <description>Correct Patern 4</description>
        <answertest>AlgEquiv</answertest>
        <sans>flgCL[4]</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>6</truenextnode>
        <trueanswernote>prt1-6-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>6</falsenextnode>
        <falseanswernote>prt1-6-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>6</name>
        <description>Correct Pattern 5</description>
        <answertest>AlgEquiv</answertest>
        <sans>flgCL[5]</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>7</truenextnode>
        <trueanswernote>prt1-7-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>7</falsenextnode>
        <falseanswernote>prt1-7-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>7</name>
        <description>Wrong(Incorrect) Pattern 1</description>
        <answertest>AlgEquiv</answertest>
        <sans>flgWL[1]</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>8</truenextnode>
        <trueanswernote>prt1-8-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>8</falsenextnode>
        <falseanswernote>prt1-8-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>8</name>
        <description>Wrong(Incorrect) Pattern 2</description>
        <answertest>AlgEquiv</answertest>
        <sans>flgWL[2]</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>9</truenextnode>
        <trueanswernote>prt1-9-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>9</falsenextnode>
        <falseanswernote>prt1-9-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>9</name>
        <description>Wrong(Incorrect) Pattern 3</description>
        <answertest>AlgEquiv</answertest>
        <sans>flgWL[3]</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>10</truenextnode>
        <trueanswernote>prt1-10-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>10</falsenextnode>
        <falseanswernote>prt1-10-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>
