<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 115594  -->
  <question type="stack">
    <name>
      <text>005.MapInv</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p><br></p>
<div>集合\(A={@SA@}\),\(B={@SB@}\)に対し,</div>
<div>写像\(f:A\to B\)を次の対応(\(A\times B\)の部分集合)\(F\)で定める.</div>
<div>&nbsp; &nbsp; &nbsp; &nbsp; &nbsp; &nbsp;\(F=\{{@SF@}\}\)</div>
<div>このとき, 逆像 \(S=f^{-1}({@xb@})\)を求めよ. 解答形式は, 答が\(\{1,2,3\}\)ならば [1,2,3], 答が \(\{1\}\) ならば [1] のようにリスト形式で答えること.</div>
<p>\(S=\)[[input:ans1]] [[validation:ans1]]</p>
<p>また, 逆像 \(f^{-1}(b)\)が空集合となるような \(b\in B\)を1つ答えよ. 
[[input:ans2]][[validation:ans2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[na:rand([6,7,7,7,8,8]);
LA:makelist(i,i,1,na);
nb:na+rand([-1,0,1,1,2,3]);
nb1:nb-rand([2,3,3,4]);
LB:makelist(i,i,1,nb);
LBr:random_permutation(LB);
LB1:makelist(LBr[i],i,1,nb1);
SA:setify(LA);
SB:setify(LB);
FixP(s):=ssubst("(","[",ssubst(")","]",s));
LF:makelist([LA[i],rand(LB1)],i,1,na);
SF:FixP(string(setify(LF)));
ImF:makelist(LF[i][2],i,1,length(LF));
xb:rand(ImF);
ImFout:listify(setdifference(SB,setify(ImF)));
mkCans(LF,y):=block([CL,i],CL:[],for i:1 thru na do if member([LA[i],y],LF) then CL:append(CL,[LA[i]]),return(CL));
Cans1:sort(mkCans(LF,xb));
Cans2:rand(ImFout);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@LA@},{@LB@},{@SF@},{@xb@},{@Cans1@},{@ImF@},{@ImFout@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>よくできました。正解です!</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>惜しい！部分的に正解です。</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>残念，間違いです。</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>(</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>Cans1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>[?,?]</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>Cans2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>if not listp(ans1) then Tans1:[ans1] else Tans1:ans1;
Sans1:sort(Tans1);
chk1:is(equal(Sans1,Cans1));
chk2:is(member(ans2,ImFout));
chk3:is(subsetp(Sans1,Cans1));</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>chk1</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">逆像\(S\)が違います.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>chk3</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty></truepenalty>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">逆像は正しくありませんが, 何か見落としがあるようです.<br></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">逆像の定義を確認してみよう.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>chk2</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">{@ans2@}の逆像は{@setify(mkCans(LF,ans2))@}で空集合ではありません.<br></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>262066930</deployedseed>
    <deployedseed>35345566</deployedseed>
    <deployedseed>596781430</deployedseed>
    <deployedseed>1209874510</deployedseed>
    <deployedseed>1855338945</deployedseed>
    <deployedseed>570731334</deployedseed>
    <deployedseed>1393577963</deployedseed>
    <deployedseed>148688295</deployedseed>
    <deployedseed>903246903</deployedseed>
    <deployedseed>1515543745</deployedseed>
  </question>

</quiz>