<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 118206  -->
  <question type="stack">
    <name>
      <text>005.DetDefDet3Prop-ElementaryColOp-1_ja</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p></p>
<p>3次正方行列の行列式 \({@dA@}\) が \(k\)であるとき,次の行列式の値は\(k\)の何倍か.</p>
<p>数値は, \(-\dfrac35\)ならば -3/5 のように 半角数字記号で解答すること.</p>
<div>(1) \({@mondai1@}=\)[[input:ans1]] [[validation:ans1]] \(k\)</div>
<div>(2) \({@mondai2@}=\)[[input:ans2]] [[validation:ans2]] \(k\)</div>
<div>(3) \({@mondai3@}=\)[[input:ans3]] [[validation:ans3]] \(k\)</div>
<p><br></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* 3次行列式の性質と列基本変形 */
/* v004 : 問題文のミスを修正 (2023.6.13) */
stack_include("https://stack.mathedu.jp/sc/ky_linear_algebra.mac");
sz:3;
n:sz;
diag_mat(L):=block([l], l:length(L), return(genmatrix(lambda([i, j], if i=j then L[i] else 0), l, l)));
matA:matrix(makelist(Wa[i], i, 1, sz));
MtoTeX(M):=block([m, n, i, j, s], s:"\\begin{vmatrix}", [m, n]:matrix_size(M), for i:1 thru m do (for j:1 thru n do (s:sconcat(s, M[i, j]), if j<n then s:sconcat(s, "&")), if i<m then s:sconcat(s, "\\\\")), return(sconcat(s, "\\end{vmatrix}")));
aconv(s):=ssubst("a_1", "a1", ssubst("a_2", "a2", ssubst("a_3", "a3", ssubst("", "*", ssubst("", "(", ssubst("", ")", s))))));
vmat(m):=block([_lc:lmxchar, s], lmxchar:"|", s:tex1(m), lmxchar:_lc, return(s));

dA:vmat(matA);
aL:[-3, -2, -1, 1, 2, 3, 4, 5];
cL:[-2, -1, 2, 3];
X1:rand([kihonP(n, 1, 2, rand(aL)), kihonP(n, 1, 3, rand(aL)), kihonP(n, 2, 3, rand(aL))]);
X2:rand([kihonQ(n, 1, rand(cL)), kihonQ(n, 2, rand(cL)).kihonP(n, 1, 2, rand(aL)), kihonQ(n, 2, rand(cL)).kihonQ(n, 3, rand(cL))]);
X3:rand([diag_mat([rand(cL), rand(cL), rand(cL)]), kihonR(n, 1, 2).kihonR(n, 2, 3).kihonQ(n, 2, rand(cL))]);
Cans1:determinant(X1);
Cans2:determinant(X2);
Cans3:determinant(X3);
CansL:[Cans1, Cans2, Cans3];
m1:matA.X1;
m2:matA.X2;
m3:matA.X3;
mondai1:vmat(m1);
mondai2:vmat(m2);
mondai3:vmat(m3);
]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@X1@},{@X2@},{@X3@}: {@m1@},{@m2@},{@m3@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>よくできました。正解です!</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>惜しい！部分的に正解です。</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>残念，間違いです。</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>(</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>Cans1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>Cans2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans3</name>
      <type>numerical</type>
      <tans>Cans3</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>ansL:[ans1,ans2,ans3];</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>Cans1</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.33</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">(1) 正解です.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>Cans2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.33</truescore>
        <truepenalty></truepenalty>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">(2) 正解です.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>Cans3</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.34</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">(3)正解です.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1509275016</deployedseed>
    <deployedseed>1144512936</deployedseed>
    <deployedseed>1970877195</deployedseed>
    <deployedseed>655192193</deployedseed>
    <deployedseed>1022651934</deployedseed>
    <deployedseed>2069898080</deployedseed>
    <deployedseed>2054699367</deployedseed>
    <deployedseed>859819339</deployedseed>
    <deployedseed>1619548418</deployedseed>
    <deployedseed>1203371450</deployedseed>
    <deployedseed>1774795114</deployedseed>
    <deployedseed>1370281166</deployedseed>
    <deployedseed>5017366</deployedseed>
    <deployedseed>262556676</deployedseed>
    <deployedseed>1378408406</deployedseed>
    <deployedseed>4443529</deployedseed>
    <deployedseed>1803682953</deployedseed>
    <deployedseed>1986177627</deployedseed>
    <deployedseed>656118630</deployedseed>
    <deployedseed>582182555</deployedseed>
  </question>

</quiz>