<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 116242  -->
  <question type="stack">
    <name>
      <text>005.MatProdRowColumnA-easy5</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p><br></p>
<p>次の計算をしなさい.</p>
<p><br></p>
<div>(1)&nbsp;&nbsp;{@rvL[1]@}{@cvL[1]@}\(=\)[[input:ans1]][[validation:ans1]]</div>
<div>(2)&nbsp;&nbsp;{@rvL[2]@}{@cvL[2]@}\(=\)[[input:ans2]][[validation:ans2]]</div>
<div>(3)&nbsp;&nbsp;{@rvL[3]@}{@cvL[3]@}\(=\)[[input:ans3]][[validation:ans3]]</div>
<div>(4)&nbsp;&nbsp;{@rvL[4]@}{@cvL[4]@}\(=\)[[input:ans4]][[validation:ans4]]</div>
<div>(5)&nbsp;&nbsp;{@rvL[5]@}{@cvL[5]@}\(=\)[[input:ans5]][[validation:ans5]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* number of Questions */
Qn:10;
matE2:genmatrix(lambda([i,j],if i=j then 1 else 0),2,2);
matE3:genmatrix(lambda([i,j],if i=j then 1 else 0),3,3);
matE4:genmatrix(lambda([i,j],if i=j then 1 else 0),4,4);
matE5:genmatrix(lambda([i,j],if i=j then 1 else 0),5,5);
matE6:genmatrix(lambda([i,j],if i=j then 1 else 0),6,6);
_rp(L):=random_permutation(L);
Qn:5;/* 5-10 */
/* list of sizes */
nL:[2,2,3,4,5,3,4,4,5,6];
nLL:makelist(_rp(makelist(j,j,1,nL[k])),k,1,10);
eL:makelist([nLL[k][1],nLL[k][2]],k,1,10);
ev2L:makelist(col(matE2,k),k,1,2);
ev3L:makelist(col(matE3,k),k,1,3);
ev4L:makelist(col(matE4,k),k,1,4);
ev5L:makelist(col(matE5,k),k,1,5);
ev6L:makelist(col(matE6,k),k,1,6);
evLL:[ev2L,ev3L,ev4L,ev5L,ev6L];
rvL:map(transpose,makelist(rand([-1,1])*evLL[nL[k]-1][eL[k][1]]-rand([-1,1])*evLL[nL[k]-1][eL[k][2]],k,1,10));
cvL:makelist(genmatrix(lambda([i,j],rand([1,2,3,4])),nL[k],1),k,1,10);
CaL:makelist(rvL[k].cvL[k],k,1,10);
mkta(C,m,M):=makelist([k-m+2,if k=C then true else false,if k=m-1 then "-" else k],k,m-1,M);
taL:makelist(mkta(CaL[k],-8,8),k,1,10);
CansL:flatten(makelist(mcq_correct(taL[k]),k,1,10));]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="html">
      <text>{@nL@},{@nLL@},{@rvL@},{@cvL@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>よくできました。正解です!</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>惜しい！部分的に正解です。</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>残念，間違いです。</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>(</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>taL[1]</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered,hideanswer</options>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>taL[2]</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered,hideanswer</options>
    </input>
    <input>
      <name>ans3</name>
      <type>dropdown</type>
      <tans>taL[3]</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered,hideanswer</options>
    </input>
    <input>
      <name>ans4</name>
      <type>dropdown</type>
      <tans>taL[4]</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered,hideanswer</options>
    </input>
    <input>
      <name>ans5</name>
      <type>dropdown</type>
      <tans>taL[5]</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered,hideanswer</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>ansL:makelist([ans1,ans2,ans3,ans4,ans5,ans6,ans7,ans8,ans9,ans10][k],k,1,Qn);
chkL(L1,L2):=block([r:0,k],for k:1 thru length(L1) do if L1[k]=L2[k] then r:r+1,return(r));
cr:chkL(ansL,CansL);
sc:cr/Qn;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>cr</sans>
        <tans>Qn</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>sc</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text>全問正解です. 簡単ですね. {@ansL@}</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>sc</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>どれか計算ミスでしょうか? {@cr@} 個正解です.</p>
{@ansL@},{@CansL@}]]></text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <description>管理者のインプットを想定したテストケースは満点です。</description>
      <testinput>
        <name>ans1</name>
        <value>CansL[1]</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>CansL[2]</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>CansL[3]</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>CansL[4]</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>CansL[5]</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>