<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 118280  -->
  <question type="stack">
    <name>
      <text>005.DetPropMultiplicativityInvertibility-1</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p><br></p>
{@qtext@}
<div>\(p\) [[input:ans1]] [[validation:ans1]] [[input:ans2]][[validation:ans2]]</div>
[[lang code='en']][[/lang]]
[[lang code='fr']][[/lang]]
[[lang code='de']][[/lang]]
[[lang code='it']][[/lang]]
[[lang code='ja']][[/lang]]]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[stack_include("https://stack.mathedu.jp/sc/ky_linear_algebra.mac");
stack_include("https://stack.mathedu.jp/sc/multilang_library.mac");
tex2(m):=sconcat("\\(", tex1(m), "\\)");
sz:4;
aL:make_random_distribution_pm([[4, 1], [3, 2], [2, 5], [1, 20], [0, 1]]);
matP1:random_matrix_det_list(sz-1, aL, [-6, -5, -4, -3, -2, -1, 1, 2, 3, 4, 5, 6]);
%_ij:[rand(sz)+1, rand(sz)+1];
%_i:%_ij[1];
%_j:%_ij[2];
matP:genmatrix(lambda([i, j], 
  if i<%_i then (if j<%_j then matP1[i, j] else if j>%_j then matP1[i, j-1] else rand(aL))
  else if i>%_i then (if j<%_j then matP1[i-1, j] else if j>%_j then matP1[i-1, j-1] else rand(aL))
  else if not j=%_j then rand(aL) else p), sz, sz);
detP:expand(determinant(matP));

va1Y:[[0, false, "--"], [1, false, "≠"], [2, true, "="]];
va1N:[[0, false, "--"], [1, true, "≠"], [2, false, "="]];
ta2:subst(solve(detP=0), p);
remove_paren(s):=regex_subst("-\\1", "-\\\\left\\(([0-9a-zA-Z \\\\{}_]+)\\\\right\\)", s);
tex2rp(s):=remove_paren(tex2(s));

qtextY:%__lang([
["en", sconcat("The determinant of the square matrix \\(A=\\)", tex2(matP), " is ", tex2rp(detP), ". Find the condition on \\(p\\) for which \\(A\\) is <b>singular</b> (that is, not invertible).")],
["fr", sconcat("Le déterminant de la matrice carrée \\(A=\\)", tex2(matP), " est ", tex2rp(detP), ". Déterminez la condition sur \\(p\\) pour que \\(A\\) soit <b>singulière</b> (c’est-à-dire non inversible).")],
["de", sconcat("Die Determinante der quadratischen Matrix \\(A=\\)", tex2(matP), " ist ", tex2rp(detP), ". Bestimmen Sie die Bedingung an \\(p\\), damit \\(A\\) <b>singulär</b> (also nicht invertierbar) ist.")],
["it", sconcat("Il determinante della matrice quadrata \\(A=\\)", tex2(matP), " è ", tex2rp(detP), ". Determina la condizione su \\(p\\) affinché \\(A\\) sia <b>singolare</b> (cioè non invertibile).")],
["ja", sconcat("正方行列\\(A=\\)", tex2(matP), "の行列式は, ", tex2rp(detP), "である. \\(A\\)が<b>非</b>正則(正則でない)ための\\(p\\)の条件を求めよ.")]
]);
qtextN:%__lang([
["en", sconcat("The determinant of the square matrix \\(A=\\)", tex2(matP), " is ", tex2rp(detP), ". Find the condition on \\(p\\) for which \\(A\\) is <b>invertible</b>.")],
["fr", sconcat("Le déterminant de la matrice carrée \\(A=\\)", tex2(matP), " est ", tex2rp(detP), ". Déterminez la condition sur \\(p\\) pour que \\(A\\) soit <b>inversible</b>.")],
["de", sconcat("Die Determinante der quadratischen Matrix \\(A=\\)", tex2(matP), " ist ", tex2rp(detP), ". Bestimmen Sie die Bedingung an \\(p\\), damit \\(A\\) <b>invertierbar</b> ist.")],
["it", sconcat("Il determinante della matrice quadrata \\(A=\\)", tex2(matP), " è ", tex2rp(detP), ". Determina la condizione su \\(p\\) affinché \\(A\\) sia <b>invertibile</b>.")],
["ja", sconcat("正方行列\\(A=\\)", tex2(matP), "の行列式は, ", tex2rp(detP), "である. \\(A\\)が正則であるための\\(p\\)の条件を求めよ.")]
]);

%_YN:rand([1, 2]);
qtext:[qtextY, qtextN][%_YN];
va1:[va1Y, va1N][%_YN];
ta1:[2, 1][%_YN];

Msg1:%__lang([
["en", "It is easy to see that a necessary condition for a matrix to be invertible (that is, to have an inverse matrix) is that its determinant is not equal to \\(0\\), since the determinant of a product is equal to the product of the determinants. In fact, this condition is also sufficient."],
["fr", "Il est facile de voir qu’une condition nécessaire pour qu’une matrice soit inversible (c’est-à-dire qu’elle possède une matrice inverse) est que son déterminant soit différent de \\(0\\), puisque le déterminant d’un produit est égal au produit des déterminants. En fait, cette condition est également suffisante."],
["de", "Es ist leicht zu erkennen, dass eine notwendige Bedingung dafür, dass eine Matrix invertierbar ist (d. h. eine Inverse besitzt), darin besteht, dass ihre Determinante nicht gleich \\(0\\) ist, da die Determinante eines Produkts gleich dem Produkt der Determinanten ist. Tatsächlich ist diese Bedingung auch hinreichend."],
["it", "È facile vedere che una condizione necessaria affinché una matrice sia invertibile (cioè possieda una matrice inversa) è che il suo determinante sia diverso da \\(0\\), poiché il determinante di un prodotto è uguale al prodotto dei determinanti. In effetti, questa condizione è anche sufficiente."],
["ja", "行列の正則性(逆行列を持つかどうか)は, 積の行列式が行列式の積になることから, 行列式が\\(0\\)でないことが必要条件であることがすぐわかります. 実は, これが十分条件にもなっています."]
]);



]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="html">
      <text>{@matP@}, {@detP@}, {@ta1@},{@ta2@}</text>
    </questionnote>
    <questiondescription format="html">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> 
<span lang="ja" class="multilang">よくできました。正解です！</span>
<span lang="en" class="multilang">Well done! Spot on!</span>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span>
<span lang="ja" class="multilang">部分的に正解です。</span>
<span lang="en" class="multilang">Partially correct!</span>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span>
<span lang="ja" class="multilang">間違いです。</span>
<span lang="en" class="multilang">Incorrect!</span>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>(</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>va1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered,hideanswer</options>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>ta2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>allowempty,hideanswer</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>[ans1, ans2]</sans>
        <tans>[ta1, ta2]</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>OK.</p>
[[lang code='en']][[/lang]]
[[lang code='fr']][[/lang]]
[[lang code='de']][[/lang]]
[[lang code='it']][[/lang]]
[[lang code='ja']][[/lang]]]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>{@Msg1@}</p>
[[lang code='en']][[/lang]]
[[lang code='fr']][[/lang]]
[[lang code='de']][[/lang]]
[[lang code='it']][[/lang]]
[[lang code='ja']][[/lang]]]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>