<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 115592  -->
  <question type="stack">
    <name>
      <text>005.MapComposite</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p><br></p>
<p>集合\(A={@SA@},B={@SB@},C={@SC@}\)に対し, 写像\(f:A\to B\)と\(g:B\to C\)を各々次の対応(\(A\times B\)の部分集合) \(F,G\)で定める.</p><p>\(F=\{{@SF@}\}\)</p><p>\(G=\{{@SG@}\}\)</p><p>このとき, \(g\circ f({@xa@})\)を答えよ.</p>
<p>[[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[LA:[1,2,3,4];na:length(LA);
LB:[1,2,3,4,5];nb:length(LB);
LC:[1,2,3];nc:lenght(LC);
SA:setify(LA);
SB:setify(LB);
SC:setify(LC);
FixP(s):=ssubst("(","[",ssubst(")","]",s));
LF:makelist([i,rand(LB)],i,1,na);
LG:makelist([LB[i],rand(LC)],i,1,nb);
SF:FixP(string(setify(LF)));
SG:FixP(string(setify(LG)));
xa:rand(LA);
Cans1:LG[LF[xa][2]][2];]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@SF@},{@SG@},{@xa@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>よくできました。正解です!</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>惜しい！部分的に正解です。</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>残念，間違いです。</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>(</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>Cans1</tans>
      <boxsize>2</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>Cans1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>938645910</deployedseed>
    <deployedseed>1974940693</deployedseed>
    <deployedseed>1472722649</deployedseed>
    <deployedseed>1985101340</deployedseed>
    <deployedseed>1581565881</deployedseed>
    <deployedseed>216546260</deployedseed>
    <deployedseed>199590604</deployedseed>
    <deployedseed>1410092893</deployedseed>
    <deployedseed>1268182494</deployedseed>
    <deployedseed>1012834170</deployedseed>
  </question>

</quiz>