<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 118254  -->
  <question type="stack">
    <name>
      <text>005.DetDefDetOfPermMat-1</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p><br></p>
<p>{@qtext@}</p>
<p>[[input:ans1]] [[validation:ans1]]</p>
[[lang code='en']][[/lang]]
[[lang code='fr']][[/lang]]
[[lang code='de']][[/lang]]
[[lang code='it']][[/lang]]
[[lang code='ja']][[/lang]]]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[[[lang code="ja"]]

<p>
\(\def\Wa{\boldsymbol a}D(A)=D(...,a_{ij},...) = D(\Wa_1, \Wa_2, \dots, \Wa_n)\)が
多重線形性と退化性を満たすと仮定すると,
</p>

<p>
\(\def\We{\boldsymbol e}a_{\sigma_11}a_{\sigma_22}\cdots a_{\sigma_nn}D(\We_{\sigma_1},\We_{\sigma_2},\dots,\We_{\sigma_n})\)の形の項の和になるのでした.
</p>

<p>
ここで, \(\sigma=[\sigma_1,\sigma_2,\dots,\sigma_n]\)は\([1,2,\dots,n]\)の順列すべてを動くということでした.
</p>

<p>
ここで, 交代性を用いると, \(D(\We_{\sigma_1},\We_{\sigma_2},\dots,\We_{\sigma_n})\)は\(\sigma\)の反転数\(N_\sigma\)を用いて, \((-1)^{N_\sigma}=\pm 1\)となるので, これを\(a_{\sigma_11}a_{\sigma_22}\cdots a_{\sigma_nn}\)の符号数と呼ぶことにします.
</p>

<p>
ここでは, この符号数を計算してもらっているわけです.
</p>
[[/lang]]

[[lang code="en"]]

<p>
Suppose that
\(\def\Wa{\boldsymbol a}
D(A)=D(\ldots,a_{ij},\ldots)=D(\Wa_1,\Wa_2,\dots,\Wa_n)\)
satisfies multilinearity and degeneracy.
</p>

<p>
Then the determinant can be expressed as a sum of terms of the form
\(\def\We{\boldsymbol e}
a_{\sigma_11}a_{\sigma_22}\cdots a_{\sigma_nn}
D(\We_{\sigma_1},\We_{\sigma_2},\dots,\We_{\sigma_n})\).
</p>

<p>
Here,
\(\sigma=[\sigma_1,\sigma_2,\dots,\sigma_n]\)
runs through all permutations of
\([1,2,\dots,n]\).
</p>

<p>
Using the alternating property,
\(D(\We_{\sigma_1},\We_{\sigma_2},\dots,\We_{\sigma_n})=(-1)^{N_\sigma}=\pm1\),
where \(N_\sigma\) denotes the inversion number of the permutation \(\sigma\).
We call this value the sign of
\(a_{\sigma_11}a_{\sigma_22}\cdots a_{\sigma_nn}\).
</p>

<p>
In this exercise, you are asked to compute this sign.
</p>
[[/lang]]

[[lang code="fr"]]

<p>
Supposons que
\(\def\Wa{\boldsymbol a}
D(A)=D(\ldots,a_{ij},\ldots)=D(\Wa_1,\Wa_2,\dots,\Wa_n)\)
satisfasse la multilinéarité et la dégénérescence.
</p>

<p>
Alors, le déterminant peut être exprimé comme une somme de termes de la forme
\(\def\We{\boldsymbol e}
a_{\sigma_11}a_{\sigma_22}\cdots a_{\sigma_nn}
D(\We_{\sigma_1},\We_{\sigma_2},\dots,\We_{\sigma_n})\).
</p>

<p>
Ici,
\(\sigma=[\sigma_1,\sigma_2,\dots,\sigma_n]\)
parcourt toutes les permutations de
\([1,2,\dots,n]\).
</p>

<p>
En utilisant la propriété alternée,
\(D(\We_{\sigma_1},\We_{\sigma_2},\dots,\We_{\sigma_n})=(-1)^{N_\sigma}=\pm1\),
où \(N_\sigma\) désigne le nombre d’inversions de la permutation \(\sigma\).
Nous appelons cette valeur le signe de
\(a_{\sigma_11}a_{\sigma_22}\cdots a_{\sigma_nn}\).
</p>

<p>
Dans cet exercice, il vous est demandé de calculer ce signe.
</p>
[[/lang]]

[[lang code="de"]]

<p>
Angenommen,
\(\def\Wa{\boldsymbol a}
D(A)=D(\ldots,a_{ij},\ldots)=D(\Wa_1,\Wa_2,\dots,\Wa_n)\)
erfüllt die Multilinearität und die Degeneriertheit.
</p>

<p>
Dann lässt sich die Determinante als Summe von Termen der Form
\(\def\We{\boldsymbol e}
a_{\sigma_11}a_{\sigma_22}\cdots a_{\sigma_nn}
D(\We_{\sigma_1},\We_{\sigma_2},\dots,\We_{\sigma_n})\)
darstellen.
</p>

<p>
Dabei durchläuft
\(\sigma=[\sigma_1,\sigma_2,\dots,\sigma_n]\)
alle Permutationen von
\([1,2,\dots,n]\).
</p>

<p>
Unter Verwendung der alternierenden Eigenschaft gilt
\(D(\We_{\sigma_1},\We_{\sigma_2},\dots,\We_{\sigma_n})=(-1)^{N_\sigma}=\pm1\),
wobei \(N_\sigma\) die Inversionszahl der Permutation \(\sigma\) bezeichnet.
Diesen Wert nennen wir das Vorzeichen von
\(a_{\sigma_11}a_{\sigma_22}\cdots a_{\sigma_nn}\).
</p>

<p>
In dieser Aufgabe sollen Sie dieses Vorzeichen bestimmen.
</p>
[[/lang]]

[[lang code="it"]]

<p>
Supponiamo che
\(\def\Wa{\boldsymbol a}
D(A)=D(\ldots,a_{ij},\ldots)=D(\Wa_1,\Wa_2,\dots,\Wa_n)\)
soddisfi la multilinearità e la degenerazione.
</p>

<p>
Allora il determinante può essere espresso come una somma di termini della forma
\(\def\We{\boldsymbol e}
a_{\sigma_11}a_{\sigma_22}\cdots a_{\sigma_nn}
D(\We_{\sigma_1},\We_{\sigma_2},\dots,\We_{\sigma_n})\).
</p>

<p>
Qui,
\(\sigma=[\sigma_1,\sigma_2,\dots,\sigma_n]\)
varia su tutte le permutazioni di
\([1,2,\dots,n]\).
</p>

<p>
Utilizzando la proprietà alternante,
\(D(\We_{\sigma_1},\We_{\sigma_2},\dots,\We_{\sigma_n})=(-1)^{N_\sigma}=\pm1\),
dove \(N_\sigma\) indica il numero di inversioni della permutazione \(\sigma\).
Chiamiamo questo valore il segno di
\(a_{\sigma_11}a_{\sigma_22}\cdots a_{\sigma_nn}\).
</p>

<p>
In questo esercizio, ti viene chiesto di calcolare questo segno.
</p>
[[/lang]]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.5</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* 符号数の計算 AI 耐性用1題構成 */
/* v005 : 全体フィードバック (2023.6.19) */
stack_include("https://stack.mathedu.jp/sc/ky_linear_algebra.mac");
stack_include("https://stack.mathedu.jp/sc/multilang_library.mac");
sz:rand([4, 5, 6, 7]);
iv_min:[3, 4, 5, 6][sz-3];
%_iv(L):=inversion_number(L);
%_rp(L):=random_permutation(L);
nL:makelist(k, k, sz);
pL:%_rp(nL);
while (%_iv(pL)<iv_min) do pL:%_rp(nL);
iv:%_iv(pL);
ta1:((-1)^iv+3)/2;

makeQS(sz, p):=block([s, k], s:"", for k:1 thru sz do (s:sconcat(s, tex1(Wa[p[k]*10+k]))), return (s));
/* QSL:makelist(makeQS(szL[i], pnL[i]), i, 1, 4); */
QS:makeQS(sz, pL);

va1:makelist([k, if k=ta1 then true else false, if k=0 then "--" else [-1, 1][k]], k, 0, 2);
szS:sconcat("\\(", sz, "\\)");
qtext:%__lang([
["en", sconcat("When expanding the determinant of the square matrix of order ", szS, " \\(A=(a_{ij})\\), as a polynomial in its entries, determine the sign of the following term: \\(", QS, "\\)")],
["fr", sconcat("Lors du développement du déterminant de la matrice carrée d’ordre ", szS, ", \\(A=(a_{ij})\\), comme polynôme en ses coefficients, déterminez le signe du terme suivant : \\(", QS, "\\)")],
["de", sconcat("Beim Entwickeln der Determinante der quadratischen Matrix ", szS, "-ter Ordnung \\(A=(a_{ij})\\) als Polynom in ihren Einträgen, bestimmen Sie das Vorzeichen des folgenden Terms: \\(", QS, "\\)")],
["it", sconcat("Nello sviluppo del determinante della matrice quadrata di ordine ", szS, ", \\(A=(a_{ij})\\), come polinomio nelle sue componenti, determina il segno del seguente termine: \\(", QS, "\\)")],
["ja", sconcat(szS, "次正方行列\\(A=(a_{ij})\\)の行列式を成分に関する多項式として展開するとき, 次の項の符号を答えよ: \\(", QS, "\\)")]]);


]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@sz@}, {@pL@},{@iv@}, {@ta1@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>よくできました。正解です!</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>惜しい！部分的に正解です。</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>残念，間違いです。</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>va1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered,hideanswer</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <description>管理者のインプットを想定したテストケースは満点です。</description>
      <testinput>
        <name>ans1</name>
        <value>first(mcq_correct(va1))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>