<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 115786  -->
  <question type="stack">
    <name>
      <text>005.MatProdRowColumnEij-x10</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p><br></p>
 <p>次の計算をしなさい.</p>
<p><br></p>
<div>(1)&nbsp;&nbsp;{@rvL[1]@}{@cvL[1]@}\(=\)[[input:ans1]][[validation:ans1]]</div>
<div>(2)&nbsp;&nbsp;{@rvL[2]@}{@cvL[2]@}\(=\)[[input:ans2]][[validation:ans2]]</div>
<div>(3)&nbsp;&nbsp;{@rvL[3]@}{@cvL[3]@}\(=\)[[input:ans3]][[validation:ans3]]</div>
<div>(4)&nbsp;&nbsp;{@rvL[4]@}{@cvL[4]@}\(=\)[[input:ans4]][[validation:ans4]]</div>
<div>(5)&nbsp;&nbsp;{@rvL[5]@}{@cvL[5]@}\(=\)[[input:ans5]][[validation:ans5]]</div>
<div>(6)&nbsp;&nbsp;{@rvL[6]@}{@cvL[6]@}\(=\)[[input:ans6]][[validation:ans6]]</div>
<div>(7)&nbsp;&nbsp;{@rvL[7]@}{@cvL[7]@}\(=\)[[input:ans7]][[validation:ans7]]</div>
<div>(8)&nbsp;&nbsp;{@rvL[8]@}{@cvL[8]@}\(=\)[[input:ans8]][[validation:ans8]]</div>
<div>(9)&nbsp;&nbsp;{@rvL[9]@}{@cvL[9]@}\(=\)[[input:ans9]][[validation:ans9]]</div>
<div>(10)&nbsp;&nbsp;{@rvL[10]@}{@cvL[10]@}\(=\)[[input:ans10]][[validation:ans10]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text>行と列の積は、内積と同様の計算です。逆にあとで一般の数ベクトル\( u, v \) の内積は \( u \) の転置(縦横をいれかえたもの)と \( v \)の積として定義されます(ですから、この場合の計算を内積と読んでいる教科書もあります)。

一方の成分に 1 と  -1 が1つずつある例では、対応する成分の差を計算していることになります。また、一方の成分に 1 が丁度2つであとは0の場合は、対応する成分の和を計算することになります。</text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2024111900</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* number of Questions */
Qn:10;
matE2:genmatrix(lambda([i,j],if i=j then 1 else 0),2,2);
matE3:genmatrix(lambda([i,j],if i=j then 1 else 0),3,3);
matE4:genmatrix(lambda([i,j],if i=j then 1 else 0),4,4);
matE5:genmatrix(lambda([i,j],if i=j then 1 else 0),5,5);
_rp(L):=random_permutation(L);
/* list of sizes */
nL:sort(makelist(rand([2,2,3,3,3,4]),k,1,10));
nLL:makelist(_rp(makelist(j,j,1,nL[k])),k,1,10);
eL:makelist([nLL[k][1],nLL[k][2]],k,1,10);
ev2L:makelist(col(matE2,k),k,1,2);
ev3L:makelist(col(matE3,k),k,1,3);
ev4L:makelist(col(matE4,k),k,1,4);
ev5L:makelist(col(matE5,k),k,1,5);
evLL:[ev2L,ev3L,ev4L,ev5L];
cvL:makelist(rand([-1,1])*evLL[nL[k]-1][eL[k][1]]+rand([-1,1])*evLL[nL[k]-1][eL[k][2]],k,1,10);
rvL:map(transpose,makelist(genmatrix(lambda([i,j],rand([1,2,3,4])),nL[k],1),k,1,10));
CaL:makelist(rvL[k].cvL[k],k,1,10);
mkta(C,m,M):=makelist([k-m+2,if k=C then true else false,if k=m-1 then "-" else k],k,m-1,M);
taL:makelist(mkta(CaL[k],-8,8),k,1,10);
CansL:flatten(makelist(mcq_correct(taL[k]),k,1,10));]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="html">
      <text>{@nL@},{@nLL@},{@rvL@},{@cvL@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>よくできました。正解です!</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>惜しい！部分的に正解です。</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>残念，間違いです。</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>(</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>taL[1]</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered,hideanswer</options>
    </input>
    <input>
      <name>ans10</name>
      <type>dropdown</type>
      <tans>taL[10]</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered,hideanswer</options>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>taL[2]</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered,hideanswer</options>
    </input>
    <input>
      <name>ans3</name>
      <type>dropdown</type>
      <tans>taL[3]</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered,hideanswer</options>
    </input>
    <input>
      <name>ans4</name>
      <type>dropdown</type>
      <tans>taL[4]</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered,hideanswer</options>
    </input>
    <input>
      <name>ans5</name>
      <type>dropdown</type>
      <tans>taL[5]</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered,hideanswer</options>
    </input>
    <input>
      <name>ans6</name>
      <type>dropdown</type>
      <tans>taL[6]</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered,hideanswer</options>
    </input>
    <input>
      <name>ans7</name>
      <type>dropdown</type>
      <tans>taL[7]</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered,hideanswer</options>
    </input>
    <input>
      <name>ans8</name>
      <type>dropdown</type>
      <tans>taL[8]</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered,hideanswer</options>
    </input>
    <input>
      <name>ans9</name>
      <type>dropdown</type>
      <tans>taL[9]</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered,hideanswer</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>ansL:[ans1,ans2,ans3,ans4,ans5,ans6,ans7,ans8,ans9,ans10];
chkL(L1,L2):=block([r:0,k],for k:1 thru length(L1) do if L1[k]=L2[k] then r:r+1,return(r));
cr:chkL(ansL,CansL);
sc:cr/10;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>cr</sans>
        <tans>10</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>sc</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text>全問正解です. 計算は簡単ですが、何か気付きましたか?</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>sc</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>どれか計算ミスでしょうか? {@cr@} 個正解です.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1693161104</deployedseed>
    <deployedseed>1706923099</deployedseed>
    <deployedseed>574600248</deployedseed>
    <deployedseed>53343019</deployedseed>
    <deployedseed>1486190864</deployedseed>
    <deployedseed>123582397</deployedseed>
    <deployedseed>855083515</deployedseed>
    <deployedseed>796213214</deployedseed>
    <deployedseed>1622918606</deployedseed>
    <deployedseed>1413487201</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>管理者のインプットを想定したテストケースは満点です。</description>
      <testinput>
        <name>ans1</name>
        <value>CansL[1]</value>
      </testinput>
      <testinput>
        <name>ans10</name>
        <value>CansL[10]</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>CansL[2]</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>CansL[3]</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>CansL[4]</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>CansL[5]</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>CansL[6]</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>CansL[7]</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>CansL[8]</value>
      </testinput>
      <testinput>
        <name>ans9</name>
        <value>CansL[9]</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>
