<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 115656  -->
  <question type="stack">
    <name>
      <text>005.SpacePlaneParaPt</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p><br></p>
<p>3点 A {@pointA@}, B {@pointB@}, C {@pointC@}を通る平面のベクトル方程式(パラメータ表示)を求めよ.</p>
<table><thead>
<tr><th scope="col">\( \begin{pmatrix}x\\y\\z\end{pmatrix}=\)</th><th scope="col">[[input:ans1]]</th><th scope="col">\( + s\)</th><th scope="col">[[input:ans2]]</th><th scope="col">\( + t\)</th><th scope="col">[[input:ans3]]</th><th scope="col">\( \quad s,t \in {\mathbb R} \)</th></tr></thead><tbody></tbody></table>
<br><br><p></p>
<p>[[validation:ans1]][[validation:ans2]][[validation:ans3]]</p><p><br></p>
<script>
  var divs = document.getElementsByClassName('varmatrixinput');
  for(var i=0; i < divs.length; i++) { 
    divs[i].style.height = '100px';
  }
</script>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p dir="ltr" style="text-align: left;">&nbsp;パラメータ表示\({\boldsymbol x}=\){@_tmpEq@}において, \({\boldsymbol u}\)や\({\boldsymbol v}\)は平面に平行な, ゼロでない, 互いに平行でないベクトルである必要があります. また, \({\boldsymbol a}\)はこの平面上の点なので, 3点を通る条件が与えられている場合, 3点のいずれかをとればよいでしょう(もちろん, 他の任意の平面上の点でも構いません).<br></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Parametic representation of a plan given by three points */
stack_include("https://stack.mathedu.jp/sc/ky_linear_algebra.mac");
stack_include("https://stack.mathedu.jp/sc/tex_library.mac");
pL:make_random_distribution_pm([[2, 1], [1, 8], [0, 2]]);
matP:random_invertible_matrix(3, pL);
vecA:col(matP, 1);
vecB:col(matP, 2);
vecC:col(matP, 3);
vec_to_point(vec):=sconcat("\\((", tex1(sequenceify(list_matrix_entries(vec))), ")\\)");
pointA:vec_to_point(vecA);
pointB:vec_to_point(vecB);
pointC:vec_to_point(vecC);

Cans1:vecA;
Cans2:vecB-vecA;
Cans3:vecC-vecA;

vec1:vecB-vecA;
vec2:vecC-vecA;
zerovec:matrix([0], [0], [0]);

/* check error pattern of the answer(s) */
max_num_nodes:10;
num_nodes:5;
/* true: correct/false: incorrect */
check_ans(a1, a2, a3):=block([flg, ret:makelist(false, max_num_nodes)],
  ret[1]:is(rank(addcol(a1-vecA, vec1, vec2))=2),
  ret[2]:not(a2=zerovec) and not(a3=zerovec),
  ret[3]:ret[2] and is(rank(addcol(vec1, vec2, a2))=2),
  ret[4]:ret[2] and is(rank(addcol(vec1, vec2, a3))=2),
  ret[5]:is(rank(addcol(a2, a3))=2),
  ret[max_num_nodes]:apply("and", makelist(ret[k], k, 1, 5)),/* completely correct ? */
  return(ret));
MsgL:makelist("", k, 1, 10);
_tmpEq:tex2L([Wa, "+s", Wu, "+t", Wv]);
MsgL[1]:sconcat(_tmpEq,"の", tex2(Wa), "が平面上にありません.");
MsgL[2]:sconcat(_tmpEq, "の", tex2(Wu), ", ", tex2(Wv), "はゼロベクトルではいけません.");
MsgL[3]:sconcat(_tmpEq, "の", tex2(Wu), "が平面と平行ではありません.");
MsgL[4]:sconcat(_tmpEq, "の", tex2(Wv), "が平面と平行ではありません.");
MsgL[5]:sconcat(_tmpEq, "の", tex2(Wu), "と", tex2(Wv), "が平行になっています.");
MsgL[10]:"正解です. 3点があたえられた場合のパラメータ表示を正しく理解しています.";
ScL:[0.2, 0.2, 0.2, 0.2, 0.2, 0, 0, 0, 0, 0];






  

]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@matP@},{@Cans1@},{@Cans2@},{@Cans3@}
MsgL:{@MsgL@}, ScL:{@ScL@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>よくできました。正解です!</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>惜しい！部分的に正解です。</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>残念，間違いです。</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>symi</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>(</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>varmatrix</type>
      <tans>Cans1</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>[[BASIC-ALGEBRA]],[[BASIC-CALCULUS]],[[BASIC-MATRIX]]</forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>hideanswer</options>
    </input>
    <input>
      <name>ans2</name>
      <type>varmatrix</type>
      <tans>Cans2</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>[[BASIC-ALGEBRA]],[[BASIC-CALCULUS]],[[BASIC-MATRIX]]</forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>hideanswer</options>
    </input>
    <input>
      <name>ans3</name>
      <type>varmatrix</type>
      <tans>Cans3</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>[[BASIC-ALGEBRA]],[[BASIC-CALCULUS]],[[BASIC-MATRIX]]</forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>hideanswer</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>chkL:check_ans(ans1, ans2, ans3);
Sc:0;
if chkL[max_num_nodes] then Sc:1 else for k:1 thru num_nodes do Sc:Sc+(if chkL[k] then ScL[k] else 0);
</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>a in on the plane?</description>
        <answertest>AlgEquiv</answertest>
        <sans>chkL[1]</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>ScL[1]</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>{@MsgL[1]@}</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>u, v should not be 0 </description>
        <answertest>AlgEquiv</answertest>
        <sans>chkL[2]</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>ScL[2]</truescore>
        <truepenalty></truepenalty>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>{@MsgL[2]@}</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>u should be parallel to the plane</description>
        <answertest>AlgEquiv</answertest>
        <sans>chkL[3]</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>ScL[3]</truescore>
        <truepenalty></truepenalty>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>{@MsgL[3]@}</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>v should be parallel to the plane</description>
        <answertest>AlgEquiv</answertest>
        <sans>chkL[4]</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>ScL[4]</truescore>
        <truepenalty></truepenalty>
        <truenextnode>4</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>{@MsgL[4]@}</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>u, v should not be parallel</description>
        <answertest>AlgEquiv</answertest>
        <sans>chkL[5]</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>ScL[5]</truescore>
        <truepenalty></truepenalty>
        <truenextnode>5</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>5</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>{@MsgL[5]@}</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>5</name>
        <description>__dummy__</description>
        <answertest>AlgEquiv</answertest>
        <sans>chkL[6]</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>ScL[6]</truescore>
        <truepenalty></truepenalty>
        <truenextnode>6</truenextnode>
        <trueanswernote>prt1-6-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>6</falsenextnode>
        <falseanswernote>prt1-6-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>{@MsgL[6]@}</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>6</name>
        <description>__dummy__</description>
        <answertest>AlgEquiv</answertest>
        <sans>chkL[7]</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>ScL[7]</truescore>
        <truepenalty></truepenalty>
        <truenextnode>7</truenextnode>
        <trueanswernote>prt1-7-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>7</falsenextnode>
        <falseanswernote>prt1-7-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>{@MsgL[7]@}</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>7</name>
        <description>__dummy__</description>
        <answertest>AlgEquiv</answertest>
        <sans>chkL[8]</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>ScL[8]</truescore>
        <truepenalty></truepenalty>
        <truenextnode>8</truenextnode>
        <trueanswernote>prt1-8-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>8</falsenextnode>
        <falseanswernote>prt1-8-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>{@MsgL[8]@}</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>8</name>
        <description>__dummy__</description>
        <answertest>AlgEquiv</answertest>
        <sans>chkL[9]</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>ScL[9]</truescore>
        <truepenalty></truepenalty>
        <truenextnode>9</truenextnode>
        <trueanswernote>prt1-9-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>9</falsenextnode>
        <falseanswernote>prt1-9-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>{@MsgL[9]@}</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>9</name>
        <description>Final Message : Perfectly correct or not(score)</description>
        <answertest>AlgEquiv</answertest>
        <sans>chkL[10]</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-10-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>{@MsgL[10]@}</p>]]></text>
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-10-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your Score is : {@Sc@}</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>743422283</deployedseed>
    <deployedseed>717265260</deployedseed>
    <deployedseed>808638021</deployedseed>
    <deployedseed>1435666896</deployedseed>
    <deployedseed>1010645643</deployedseed>
    <deployedseed>1106470660</deployedseed>
    <deployedseed>1603670753</deployedseed>
    <deployedseed>1839219477</deployedseed>
    <deployedseed>1370547313</deployedseed>
    <deployedseed>889525376</deployedseed>
    <deployedseed>356632195</deployedseed>
    <deployedseed>128915117</deployedseed>
    <deployedseed>1664179434</deployedseed>
    <deployedseed>1521950369</deployedseed>
    <deployedseed>890700923</deployedseed>
    <deployedseed>1223526005</deployedseed>
    <deployedseed>808196429</deployedseed>
    <deployedseed>1140750691</deployedseed>
    <deployedseed>1922544386</deployedseed>
    <deployedseed>1176240923</deployedseed>
    <deployedseed>1782571264</deployedseed>
    <deployedseed>193355708</deployedseed>
    <deployedseed>223619430</deployedseed>
    <deployedseed>1060054681</deployedseed>
    <deployedseed>1664699200</deployedseed>
    <deployedseed>1389900011</deployedseed>
    <deployedseed>494348955</deployedseed>
    <deployedseed>18416105</deployedseed>
    <deployedseed>1711062102</deployedseed>
    <deployedseed>851225372</deployedseed>
  </question>

</quiz>