<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 116011  -->
  <question type="stack">
    <name>
      <text>010.MatProdMatVecEqLinComb</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p><br></p>
<p>\(A=\){@matA@}と\({@Wp@}={@vP@}\)との積 \(A{@Wp@}\)を考えよう.</p>
<p>一般に, \(X,Y\)が \(AX,AY\)の定義される同じ型の行列(またはベクトル)のとき,
 \(A(cX+dY)=cAX+dAY\)が成り立つことがわかる(分配法則). また, {@Wp@}は</p>
<p>\({@Wp@}={@p[1]@}{@We[1]@}+{@p[2]@}{@We[2]@}+\cdots\)</p>
<p>と書けるので, </p>
<p>\(A{@Wp@}={@p[1]@} A{@We[1]@}+{@p[2]@} A{@We[2]@}+\cdots\)</p>
<p>となることがわかる. すなわち, </p>
<table><tr><td>\(A{@Wp@}={@p[1]@}\)</td><td>[[input:ans1]] [[validation:ans1]]</td>
<td>\(+{@p[2]@}\)</td><td>[[input:ans2]][[validation:ans2]]</td><td>\(+\cdots\)</td></tr></table>
<p>となることがわかる. つまり行列とベクトルの積は, 行列の列ベクトルの1次結合になり, 
1次結合の係数はベクトルの成分ということになる. </p>
<script>
  var divs = document.getElementsByClassName('varmatrixinput');
  for(var i=0; i < divs.length; i++) { 
    divs[i].style.height = '125px';
  }
</script>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text>行列{@A@}とベクトル{@Wv@}の積が, 行列{@A@}の列ベクトルの{@Wv@}の成分を係数とする1次結合になるという解釈は
線形代数にとって非常に本質的です. しっかり, 理解しておきましょう.</text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/****************** 010.MatProdMatVecEqLinComb ***************************/
/* Object: To understand that the product of a matrix and a vector       */
/*          is a linear combination of the column vectors of the matrix, */
/*          with the components of the vector serving as coefficients.   */
/*************************************************************************/
stack_include("https://stack.mathedu.jp/sc/ky_linear_algebra.txt");
nmL:random_permutation(rand([[3, 4], [3, 5], [4, 5]]));
num_rows:nmL[1];num_cols:nmL[2];
i0:rand(num_cols)+1;
pL:makelist(p[k], k, 1, num_cols);
vP:transpose(matrix(pL));
j0:rand(num_rows)+1;
Wfj:genmatrix(lambda([i,j], if j=j0 then 1 else 0), 1, num_rows);
aL:make_random_distribution_pm([[3, 2], [2, 4], [1, 8], [0, 1]]);
matA:genmatrix(lambda([i,j], rand(aL)), num_rows, num_cols);
ta1:col(matA, 1);
ta2:col(matA, 2);
]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@[n,m]@},{@matA@},{@ii@},{@jj@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>よくできました。正解です!</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>惜しい！部分的に正解です。</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>残念，間違いです。</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>(</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>varmatrix</type>
      <tans>ta1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>allowempty,hideanswer</options>
    </input>
    <input>
      <name>ans2</name>
      <type>varmatrix</type>
      <tans>ta2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>allowempty, hideanswer</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text>1つ目正解.</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>1つ目が違います.</text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text>2つ目正解.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text>2つ目が違います.</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1141006913</deployedseed>
    <deployedseed>287302478</deployedseed>
    <deployedseed>465017778</deployedseed>
    <deployedseed>1025391060</deployedseed>
    <deployedseed>96877604</deployedseed>
    <deployedseed>1409195411</deployedseed>
    <deployedseed>1057682998</deployedseed>
    <deployedseed>981758443</deployedseed>
    <deployedseed>1830682255</deployedseed>
    <deployedseed>3631001</deployedseed>
    <deployedseed>2123481742</deployedseed>
    <deployedseed>355923794</deployedseed>
    <deployedseed>1875058021</deployedseed>
    <deployedseed>2039667983</deployedseed>
    <deployedseed>1184634579</deployedseed>
    <deployedseed>365856318</deployedseed>
    <deployedseed>1214352435</deployedseed>
    <deployedseed>807727360</deployedseed>
    <deployedseed>2071935089</deployedseed>
    <deployedseed>131300356</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>管理者のインプットを想定したテストケースは満点です。</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>