<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 115978  -->
  <question type="stack">
    <name>
      <text>010.MatUnitMatrixIntro</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p><br></p>
<p>次の計算をしなさい.</p>

<table>
<tr><td>(1) \({@matEn@}{@matA@}=\)</td><td>[[input:ans1]] [[validation:ans1]]</td></tr>
<tr><td>(2) \({@matA@}{@matEm@}=\)</td><td>[[input:ans2]] [[validation:ans2]]</td></tr>
</table>
<p>左辺と右辺を比べてみましょう. 何か気付いたことはありますか?</p>

<script>
  var divs = document.getElementsByClassName('varmatrixinput');
  for(var i=0; i < divs.length; i++) { 
    divs[i].style.height = '100px';
  }
</script>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text>{@matEn@}のように対角線にそった成分(\(a_{ii}\))がすべて\(1\)で他は0であるような行列を単位行列といいます. このような行列は, 行列の積において, 数の場合の\(1\)の役割を果たします.</text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/********** 010.MatUnitMatrixIntro *************/
/* Object: to learn the meaning of unit matrix */
/***********************************************/
nm:rand([[2,3],[3,2],[2,4],[3,3],[4,2],[3,4],[4,3]]);
n:nm[1];m:nm[2];
aL:[-3,-2,-2,-1,-1,-1,-1,1,1,1,1,2,2,3];
matA:genmatrix(lambda([i,j],rand(aL)),n,m);
matEn:genmatrix(lambda([i,j],if i=j then 1 else 0),n,n);
matEm:genmatrix(lambda([i,j],if i=j then 1 else 0),m,m);
ta1:matA;
ta2:matA;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@[n,m]@}, {@matA@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>よくできました。正解です!</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>惜しい！部分的に正解です。</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>残念，間違いです。</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>(</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>varmatrix</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>allowempty</options>
    </input>
    <input>
      <name>ans2</name>
      <type>varmatrix</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>allowempty</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text>(1) 正解です. 何か気付きましたか?</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>(1) 違います. もう一度計算を確認しよう.</text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text>(2) 正解です. (1)も正解であれば気付いたと思いますが, {@matEn@}や{@matEm@}はどういう性質を持つと言えますか?</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text>(2) 違います. もう一度計算を確認しよう.</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1124260533</deployedseed>
    <deployedseed>62830190</deployedseed>
    <deployedseed>1802948595</deployedseed>
    <deployedseed>294867879</deployedseed>
    <deployedseed>927609016</deployedseed>
    <deployedseed>963514542</deployedseed>
    <deployedseed>1996797839</deployedseed>
    <deployedseed>1800618009</deployedseed>
    <deployedseed>2145181590</deployedseed>
    <deployedseed>819715128</deployedseed>
    <qtest>
      <testcase>2</testcase>
      <description>管理者のインプットを想定したテストケースは満点です。</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>