<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 115984  -->
  <question type="stack">
    <name>
      <text>010.MatProdNonCommutativityA</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p><br></p>
<p>\(A={@matA@},B={@matB@}\)とする.
<p>(1) \(AB\)を求めよ. [[input:ans1]] [[validation:ans1]]</p>
<p>(2) \(BA\)を求めよ. [[input:ans2]][[validation:ans2]]</p>

<script>
  var divs = document.getElementsByClassName('varmatrixinput');
  for(var i=0; i < divs.length; i++) { 
    divs[i].style.height = '80px';
  }
</script>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>※ \(A, B\)が同じサイズの正方行列であれば, \(AB, BA\)どちらも定義されますが, \(AB=BA\)は一般に成り立ちません. </p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/************ 010.MatProdNonCommutativityA ************/
/* Object: Gain an intuitive understanding of the     */
/*         noncommutativity of matrix multiplication  */
/******************************************************/
nn:2;
aL:[-3, -2, -1, -1, 0, 1, 1, 2, 3];
bL:[1, 1, 1, 2, 2, 3, 4];
matA:genmatrix(lambda([i,j],rand(aL)), nn, nn);
matB:genmatrix(lambda([i,j],rand(bL)), nn, nn);
ta1:matA.matB;
ta2:matB.matA;
Cans1:ta1;
Cans2:ta2;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@matA@},{@matB@},{@Cans1@},{@Cans2@},{@is(Cans1=Cans2)@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>よくできました。正解です!</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>惜しい！部分的に正解です。</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>残念，間違いです。</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>(</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>varmatrix</type>
      <tans>ta1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>allowempty</options>
    </input>
    <input>
      <name>ans2</name>
      <type>varmatrix</type>
      <tans>ta2</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>allowempty</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text>(1)正解です.</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>(1)違います.</text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text>(2) 正解です.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text>(2) 違います.</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1619444347</deployedseed>
    <deployedseed>1049785711</deployedseed>
    <deployedseed>938563696</deployedseed>
    <deployedseed>1122355353</deployedseed>
    <deployedseed>93485975</deployedseed>
    <deployedseed>1473809559</deployedseed>
    <deployedseed>958913451</deployedseed>
    <deployedseed>2075154734</deployedseed>
    <deployedseed>869249822</deployedseed>
    <deployedseed>241965901</deployedseed>
    <deployedseed>471005492</deployedseed>
    <deployedseed>1933982967</deployedseed>
    <deployedseed>1616266426</deployedseed>
    <deployedseed>1582805121</deployedseed>
    <deployedseed>75401299</deployedseed>
    <deployedseed>939280314</deployedseed>
    <deployedseed>730757362</deployedseed>
    <deployedseed>1366851446</deployedseed>
    <deployedseed>1919324450</deployedseed>
    <deployedseed>1963075128</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>管理者のインプットを想定したテストケースは満点です。</description>
      <testinput>
        <name>ans1</name>
        <value>Cans1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>Cans2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>