<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 117016  -->
  <question type="stack">
    <name>
      <text>010.GaussElimEchelonFormDef</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p><br></p>[[lang code="en"]][[/lang]][[lang code="fr"]][[/lang]][[lang code="de"]][[/lang]][[lang code="it"]][[/lang]][[lang code="ja"]][[/lang]]
{@KAISETU@}
{@REIDAI@}
{@question_text@}
<div>(1) {@matAL[1]@} [[input:ans1]][[validation:ans1]] [[adapt id='1' hidden='true']]{@MsgL[1]@}[[/adapt]]</div>
<div>(2) {@matAL[2]@} [[input:ans2]][[validation:ans2]][[adapt id='2' hidden='true']]{@MsgL[2]@}[[/adapt]]</div>
<div>(3) {@matAL[3]@} [[input:ans3]][[validation:ans3]][[adapt id='3' hidden='true']]{@MsgL[3]@}[[/adapt]]</div>
<div>(4) {@matAL[4]@} [[input:ans4]][[validation:ans4]][[adapt id='4' hidden='true']]{@MsgL[4]@}[[/adapt]]</div>
<div>(5) {@matAL[5]@} [[input:ans5]][[validation:ans5]][[adapt id='5' hidden='true']]{@MsgL[5]@}[[/adapt]]</div>
<div>(6) {@matAL[6]@} [[input:ans6]][[validation:ans6]][[adapt id='6' hidden='true']]{@MsgL[6]@}[[/adapt]]</div>
<div>(7) {@matAL[7]@} [[input:ans7]][[validation:ans7]][[adapt id='7' hidden='true']]{@MsgL[7]@}[[/adapt]]</div>
<div>(8) {@matAL[8]@} [[input:ans8]][[validation:ans8]][[adapt id='8' hidden='true']]{@MsgL[8]@}[[/adapt]]</div>
<div>(9) {@matAL[9]@} [[input:ans9]][[validation:ans9]][[adapt id='9' hidden='true']]{@MsgL[9]@}[[/adapt]]</div>
<div>(10) {@matAL[10]@} [[input:ans10]][[validation:ans10]][[adapt id='10' hidden='true']]{@MsgL[10]@}[[/adapt]]</div>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Explanation of EchelonForm */
stack_include("https://stack.mathedu.jp/sc/ky_linear_algebra.mac");
stack_include("https://stack.mathedu.jp/sc/010/GaussElimEchelonFormDef.txt");
]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta1@}</text>
    </questionnote>
    <questiondescription format="html">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> 
<span lang="ja" class="multilang">よくできました。正解です！</span>
<span lang="en" class="multilang">Well done! Spot on!</span>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span>
<span lang="ja" class="multilang">部分的に正解です。</span>
<span lang="en" class="multilang">Partially correct!</span>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span>
<span lang="ja" class="multilang">間違いです。</span>
<span lang="en" class="multilang">Incorrect!</span>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>(</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>vaL[1]</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered,hideanswer</options>
    </input>
    <input>
      <name>ans10</name>
      <type>dropdown</type>
      <tans>vaL[10]</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered,hideanswer</options>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>vaL[2]</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered,hideanswer</options>
    </input>
    <input>
      <name>ans3</name>
      <type>dropdown</type>
      <tans>vaL[3]</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered,hideanswer</options>
    </input>
    <input>
      <name>ans4</name>
      <type>dropdown</type>
      <tans>vaL[4]</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered,hideanswer</options>
    </input>
    <input>
      <name>ans5</name>
      <type>dropdown</type>
      <tans>vaL[5]</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered,hideanswer</options>
    </input>
    <input>
      <name>ans6</name>
      <type>dropdown</type>
      <tans>vaL[6]</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered,hideanswer</options>
    </input>
    <input>
      <name>ans7</name>
      <type>dropdown</type>
      <tans>vaL[7]</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered,hideanswer</options>
    </input>
    <input>
      <name>ans8</name>
      <type>dropdown</type>
      <tans>vaL[8]</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered,hideanswer</options>
    </input>
    <input>
      <name>ans9</name>
      <type>dropdown</type>
      <tans>vaL[9]</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered,hideanswer</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>sc1:0.1;

</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>taL[1]</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>sc1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[[[lang code="en"]][[/lang]][[lang code="fr"]][[/lang]][[lang code="de"]][[/lang]][[lang code="it"]][[/lang]][[lang code="ja"]][[/lang]][[adaptauto show_ids='1' /]]]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[[[lang code="en"]][[/lang]][[lang code="fr"]][[/lang]][[lang code="de"]][[/lang]][[lang code="it"]][[/lang]][[lang code="ja"]][[/lang]][[adaptauto show_ids='1' /]]]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>taL[2]</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>sc1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[[[lang code="en"]][[/lang]][[lang code="fr"]][[/lang]][[lang code="de"]][[/lang]][[lang code="it"]][[/lang]][[lang code="ja"]][[/lang]][[adaptauto show_ids='2' /]]]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[[[lang code="en"]][[/lang]][[lang code="fr"]][[/lang]][[lang code="de"]][[/lang]][[lang code="it"]][[/lang]][[lang code="ja"]][[/lang]][[adaptauto show_ids='2' /]]]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>taL[3]</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>sc1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[[[lang code="en"]][[/lang]][[lang code="fr"]][[/lang]][[lang code="de"]][[/lang]][[lang code="it"]][[/lang]][[lang code="ja"]][[/lang]][[adaptauto show_ids='3' /]]]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[[[lang code="en"]][[/lang]][[lang code="fr"]][[/lang]][[lang code="de"]][[/lang]][[lang code="it"]][[/lang]][[lang code="ja"]][[/lang]][[adaptauto show_ids='3' /]]]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>taL[4]</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>sc1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>4</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[[[lang code="en"]][[/lang]][[lang code="fr"]][[/lang]][[lang code="de"]][[/lang]][[lang code="it"]][[/lang]][[lang code="ja"]][[/lang]][[adaptauto show_ids='4' /]]]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[[[lang code="en"]][[/lang]][[lang code="fr"]][[/lang]][[lang code="de"]][[/lang]][[lang code="it"]][[/lang]][[lang code="ja"]][[/lang]][[adaptauto show_ids='4' /]]]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>taL[5]</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>sc1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>5</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[[[lang code="en"]][[/lang]][[lang code="fr"]][[/lang]][[lang code="de"]][[/lang]][[lang code="it"]][[/lang]][[lang code="ja"]][[/lang]][[adaptauto show_ids='5' /]]]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>5</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[[[lang code="en"]][[/lang]][[lang code="fr"]][[/lang]][[lang code="de"]][[/lang]][[lang code="it"]][[/lang]][[lang code="ja"]][[/lang]][[adaptauto show_ids='5' /]]]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>5</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans6</sans>
        <tans>taL[6]</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>sc1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>6</truenextnode>
        <trueanswernote>prt1-6-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[[[lang code="en"]][[/lang]][[lang code="fr"]][[/lang]][[lang code="de"]][[/lang]][[lang code="it"]][[/lang]][[lang code="ja"]][[/lang]][[adaptauto show_ids='6' /]]]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>6</falsenextnode>
        <falseanswernote>prt1-6-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[[[lang code="en"]][[/lang]][[lang code="fr"]][[/lang]][[lang code="de"]][[/lang]][[lang code="it"]][[/lang]][[lang code="ja"]][[/lang]][[adaptauto show_ids='6' /]]]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>6</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans7</sans>
        <tans>taL[7]</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>sc1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>7</truenextnode>
        <trueanswernote>prt1-7-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[[[lang code="en"]][[/lang]][[lang code="fr"]][[/lang]][[lang code="de"]][[/lang]][[lang code="it"]][[/lang]][[lang code="ja"]][[/lang]][[adaptauto show_ids='7' /]]]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>7</falsenextnode>
        <falseanswernote>prt1-7-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[[[lang code="en"]][[/lang]][[lang code="fr"]][[/lang]][[lang code="de"]][[/lang]][[lang code="it"]][[/lang]][[lang code="ja"]][[/lang]][[adaptauto show_ids='7' /]]]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>7</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans8</sans>
        <tans>taL[8]</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>sc1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>8</truenextnode>
        <trueanswernote>prt1-8-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[[[lang code="en"]][[/lang]][[lang code="fr"]][[/lang]][[lang code="de"]][[/lang]][[lang code="it"]][[/lang]][[lang code="ja"]][[/lang]][[adaptauto show_ids='8' /]]]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>8</falsenextnode>
        <falseanswernote>prt1-8-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[[[lang code="en"]][[/lang]][[lang code="fr"]][[/lang]][[lang code="de"]][[/lang]][[lang code="it"]][[/lang]][[lang code="ja"]][[/lang]][[adaptauto show_ids='8' /]]]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>8</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans9</sans>
        <tans>taL[9]</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>sc1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>9</truenextnode>
        <trueanswernote>prt1-9-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[[[lang code="en"]][[/lang]][[lang code="fr"]][[/lang]][[lang code="de"]][[/lang]][[lang code="it"]][[/lang]][[lang code="ja"]][[/lang]][[adaptauto show_ids='9' /]]]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>9</falsenextnode>
        <falseanswernote>prt1-9-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[[[lang code="en"]][[/lang]][[lang code="fr"]][[/lang]][[lang code="de"]][[/lang]][[lang code="it"]][[/lang]][[lang code="ja"]][[/lang]][[adaptauto show_ids='9' /]]]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>9</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans10</sans>
        <tans>taL[10]</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>sc1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-10-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[[[lang code="en"]][[/lang]][[lang code="fr"]][[/lang]][[lang code="de"]][[/lang]][[lang code="it"]][[/lang]][[lang code="ja"]][[/lang]][[adaptauto show_ids='10' /]]]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-10-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[[[lang code="en"]][[/lang]][[lang code="fr"]][[/lang]][[lang code="de"]][[/lang]][[lang code="it"]][[/lang]][[lang code="ja"]][[/lang]][[adaptauto show_ids='10' /]]]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>