<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 115999  -->
  <question type="stack">
    <name>
      <text>010.MatProdMatBasisVectorA</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p><br></p>
<p>\(A=\){@matA@}と\({@We[i0]@}={@Wei@}\)との積 \(A{@We[i0]@}\)を求めよ.</p>
<table><tr><td>\(A{@We[i0]@}=\)</td><td>[[input:ans1]] [[validation:ans1]]</td></tr></table>
<p>このことから, 行列と基本ベクトルの積 \(A{@We[i]@}\)は\(A\)の第\(i\)[[input:ans2]][[validation:ans2]]であることがわかる.</p>
<p>同様に, \({@Wf[j0]@}={@Wfj@}\)と\(A\)との積 \({@Wf[j0]@}A\)を求めよ.</p>
<table><tr><td>\({@Wf[j0]@}A=\)</td><td>[[input:ans3]] [[validation:ans3]]</td></tr></table>
<p>このことから, 基本ベクトルの転置(行ベクトルの基本ベクトル)と行列との積 \({@Wf[j]@} A\)は\(A\)の第\(j\)[[input:ans4]][[validation:ans4]]であることがわかる.</p>
<script>
  var divs = document.getElementsByClassName('varmatrixinput');
  for(var i=0; i < divs.length; i++) { 
    divs[i].style.height = '125px';
  }
</script>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/***************** 010.MatProdMatBasisVectorA *************/
/* Object: Learn about the properties of the product of ***/
/*         a matrix and a basis vector.                   */
/**********************************************************/
stack_include("https://stack.mathedu.jp/sc/ky_linear_algebra.txt");
nmL:random_permutation(rand([[3,4],[3,5],[4,5]]));
num_rows:nmL[1];num_cols:nmL[2];
i0:rand(num_cols)+1;
Wei:genmatrix(lambda([i,j],if i=i0 then 1 else 0), num_cols,1);
j0:rand(num_rows)+1;
Wfj:genmatrix(lambda([i,j],if j=j0 then 1 else 0), 1, num_rows);
aL:make_random_distribution_pm([[3, 1], [2, 2], [1, 3], [0, 1]]);
matA:genmatrix(lambda([i,j], rand(aL)), num_rows, num_cols);
ta1:col(matA, i0);
ta2:[[0, false, "-"], [1, true, "列ベクトル"], [101, false, "行ベクトル"]];
ca2:1;
ta3:row(matA, j0) ;
ta4:[[0, false, "-"], [101, false, "列ベクトル"], [1, true, "行ベクトル"]];
ca4:1;
]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@nmL@},{@matA@},{@i0@},{@j0@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>よくできました。正解です!</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>惜しい！部分的に正解です。</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>残念，間違いです。</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>(</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>varmatrix</type>
      <tans>ta1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>allowempty, hideanswer</options>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <input>
      <name>ans3</name>
      <type>varmatrix</type>
      <tans>ta3</tans>
      <boxsize>12</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>allowempty, hideanswer</options>
    </input>
    <input>
      <name>ans4</name>
      <type>dropdown</type>
      <tans>ta4</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.3</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text>\(A{@ee@}_{@i0@}\) : 正解です.</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>\(A{@ee@}_{@i0@}\) が違います. 積の定義にたちかえってやってみよう.</text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ca2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.2</truescore>
        <truepenalty></truepenalty>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text>その通りです. 第\(i\)基本ベクトルを右からかけることは, 第\(i\)列を「とりだす」ことになるわけですね.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text>違います. 行と列を混同していませんか?</text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.3</truescore>
        <truepenalty></truepenalty>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text>\({@ff@}_{@j0@} A\) : 正解です.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text>\({@ff@}_{@j0@} A\) が違います. 積の定義にたちかえってやってみよう.</text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ca4</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.2</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text>その通りです. 行ベクトルの第\(j\)基本ベクトルを左からかけることは, 第\(j\)行を「とりだす」ことになるわけですね.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text>違います. 行と列を混同していませんか?</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1141006913</deployedseed>
    <deployedseed>287302478</deployedseed>
    <deployedseed>465017778</deployedseed>
    <deployedseed>1025391060</deployedseed>
    <deployedseed>96877604</deployedseed>
    <deployedseed>1409195411</deployedseed>
    <deployedseed>1057682998</deployedseed>
    <deployedseed>981758443</deployedseed>
    <deployedseed>1830682255</deployedseed>
    <deployedseed>3631001</deployedseed>
    <deployedseed>2123481742</deployedseed>
    <deployedseed>355923794</deployedseed>
    <deployedseed>1875058021</deployedseed>
    <deployedseed>2039667983</deployedseed>
    <deployedseed>1184634579</deployedseed>
    <deployedseed>365856318</deployedseed>
    <deployedseed>1214352435</deployedseed>
    <deployedseed>807727360</deployedseed>
    <deployedseed>2071935089</deployedseed>
    <deployedseed>131300356</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>管理者のインプットを想定したテストケースは満点です。</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>first(mcq_correct(ta2))</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>first(mcq_correct(ta4))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-4-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>管理者のインプットを想定したテストケースは満点です。</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>first(mcq_correct(ta2))</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>first(mcq_correct(ta4))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-4-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>