<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 115736  -->
  <question type="stack">
    <name>
      <text>010.MatProductA-num</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p><br></p>
{@question_sentence@}
<h2>行ベクトルと列ベクトルの積</h2>
<p>
行ベクトルと列ベクトルの積は, 対応する各要素の積をすべて足し合わせることで求められ, 1つの数値が得られます.
</p>
<p>
例として、行ベクトル {@matrix([a, b, c])@} と列ベクトル{@transpose(matrix([x, y, z]))@}の場合, 
計算は以下のようになります. 
</p>
{@matrix([a, b, c])@}{@transpose(matrix([x, y, z]))@}\(=a \times x + b \times y + c \times z \)
</p>
<p>
このように, 各要素同士の積の和を計算することで, 行ベクトルと列ベクトルの積は1つの数となります.
</p>
<p>
具体例として, 行ベクトル {@matA0@} と列ベクトル  {@matB0@}の場合, 積は
</p>
<p style="text-align: center;">
{@matA0B0S@}
</p>
<p>
となります.
</p>
<p>では, いくつか計算してみましょう.</p>
<p>(1) {@matA1@}{@matB1@}\(=\) [[input:ans1]] [[validation:ans1]]</p>
<p>(2) {@matA2@}{@matB2@}\(=\) [[input:ans2]] [[validation:ans2]]</p>
<p>(3) {@matA3@}{@matB3@}\(=\) [[input:ans3]] [[validation:ans3]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* 010.Explanation and Warm-up Questions */
stack_include("https://stack.mathedu.jp/sc/tex_library.mac");
/* stack_include("https://stack.mathedu.jp/sc/multilang_library.mac");*/

tex1sig(n):=sconcat(if n<0 then "(" else "", tex1(n), if n<0 then ")" else "");
n0:3;
cL0:random_permutation([1, 1, 2, 2, 3, 4, 5]);
matA0:genmatrix(lambda([i,j], cL0[j]), 1, n0);
matB0:genmatrix(lambda([i,j], cL0[i+3]), n0, 1);
matA0B0S:"\\(";
for k:1 thru n0 do matA0B0S:sconcat(matA0B0S, tex1(matA0[1, k]), "\\times", tex1(matB0[k, 1]), if k<n0 then "+" else "");
matA0B0S:sconcat(matA0B0S, "=");
for k:1 thru n0 do matA0B0S:sconcat(matA0B0S, tex1(matA0[1, k]*matB0[k, 1]), if k<n0 then "+" else "");
matA0B0S:sconcat(matA0B0S, "=");
matA0B0S:sconcat(matA0B0S, tex1(matA0.matB0));
matA0B0S:sconcat(matA0B0S, ".\\)");
/* example 2 */
n1:2;
cL1:random_permutation([1, 1, 2, 3, 4, 5, 6]);
matA1:genmatrix(lambda([i,j], cL1[j]), 1, n1);
matB1:genmatrix(lambda([i,j], cL1[i+n1]), n1, 1);
matA1B1S:"\\(";
for k:1 thru n1 do matA1B1S:sconcat(matA1B1S, tex1sig(matA1[1, k]), "\\times", tex1sig(matB1[k, 1]), if k<n1 then "+" else "");
matA1B1S:sconcat(matA1B1S, "=");
for k:1 thru n1 do matA1B1S:sconcat(matA1B1S, tex1sig(matA1[1, k]*matB1[k, 1]), if k<n1 then "+" else "");
matA1B1S:sconcat(matA1B1S, "=");
matA1B1S:sconcat(matA1B1S, tex1(matA1.matB1));
matA1B1S:sconcat(matA1B1S, ".\\)");
/* simpel questions */
n2:3;
cL2:random_permutation([-2, -1, 1, 2, 3, 4, 5, 6]);
matA2:genmatrix(lambda([i,j], cL2[j]), 1, n2);
matB2:genmatrix(lambda([i,j], cL2[i+n2]), n2, 1);
matA2B2S:"\\(";
for k:1 thru n2 do matA2B2S:sconcat(matA2B2S, tex1sig(matA2[1, k]), "\\times", tex1sig(matB2[k, 1]), if k<n2 then "+" else "");
matA2B2S:sconcat(matA2B2S, "=");
for k:1 thru n2 do matA2B2S:sconcat(matA2B2S, tex1sig(matA2[1, k]*matB2[k, 1]), if k<n2 then "+" else "");
matA2B2S:sconcat(matA2B2S, "=");
matA2B2S:sconcat(matA2B2S, tex1(matA2.matB2));
matA2B2S:sconcat(matA2B2S, ".\\)");
n3:4;
cL3:random_permutation([-3, -2, -1, -1, 0, 1, 1, 2, 3]);
matA3:genmatrix(lambda([i,j], cL3[j]), 1, n3);
matB3:genmatrix(lambda([i,j], cL3[i+n3]), n3, 1);
matA3B3S:"\\(";
for k:1 thru n3 do matA3B3S:sconcat(matA3B3S, tex1sig(matA3[1, k]), "\\times", tex1sig(matB3[k, 1]), if k<n3 then "+" else "");
matA3B3S:sconcat(matA3B3S, "=");
for k:1 thru n3 do matA3B3S:sconcat(matA3B3S, tex1sig(matA3[1, k]*matB3[k, 1]), if k<n3 then "+" else "");
matA3B3S:sconcat(matA3B3S, "=");
matA3B3S:sconcat(matA3B3S, tex1(matA3.matB3));
matA3B3S:sconcat(matA3B3S, ".\\)");
ta1:matA1.matB1;
ta2:matA2.matB2;
ta3:matA3.matB3;
qtextL:[["en", sconcat("<h2>Product of a Row Vector and a Column Vector</h2>",
 "<p>The product of a row vector and a column vector is obtained by summing the products of the corresponding components, resulting in a single number.</p>",
 "<p>For example, for the row vector ", tex1(matrix([a, b, c])), " and the column vector ", tex1(transpose(matrix([x, y, z]))), ", ",
 "the computation is as follows.</p><p>\\(",
 tex1(matrix([a, b, c])), tex1(transpose(matrix([x, y, z]))),
 "\\(=a \\times x + b \\times y + c \\times z \\)</p>",
 "<p>In this way, by calculating the sum of the products of corresponding components, the product of a row vector and a column vector is a single number.</p>",
 "<p>As a concrete example, for the row vector ", tex1(matA0), " and the column vector ", tex1(matB0), ", the product is</p>",
 "<p style=\"text-align: center;\">", tex1(matA0B0S), "</p><p>.</p>",
 "<p>Now, let’s try some calculations.</p>")],
["fr", sconcat("<h2>Produit d’un vecteur ligne et d’un vecteur colonne</h2>",
 "<p>Le produit d’un vecteur ligne et d’un vecteur colonne s’obtient en additionnant les produits des composantes correspondantes, ce qui donne un seul nombre.</p>",
 "<p>Par exemple, pour le vecteur ligne ", tex1(matrix([a, b, c])), " et le vecteur colonne ", tex1(transpose(matrix([x, y, z]))), ", ",
 "le calcul est le suivant.</p><p>\\(",
 tex1(matrix([a, b, c])), tex1(transpose(matrix([x, y, z]))),
 "\\(=a \\times x + b \\times y + c \\times z \\)</p>",
 "<p>Ainsi, en calculant la somme des produits des composantes correspondantes, le produit d’un vecteur ligne et d’un vecteur colonne est un seul nombre.</p>",
 "<p>Comme exemple concret, pour le vecteur ligne ", tex1(matA0), " et le vecteur colonne ", tex1(matB0), ", le produit est</p>",
 "<p style=\"text-align: center;\">", tex1(matA0B0S), "</p><p>.</p>",
 "<p>Essayons maintenant quelques calculs.</p>")],
["de", sconcat("<h2>Produkt eines Zeilenvektors und eines Spaltenvektors</h2>",
 "<p>Das Produkt eines Zeilenvektors und eines Spaltenvektors erhält man, indem man die Produkte der entsprechenden Komponenten aufsummiert, wodurch eine einzelne Zahl entsteht.</p>",
 "<p>Zum Beispiel für den Zeilenvektor ", tex1(matrix([a, b, c])), " und den Spaltenvektor ", tex1(transpose(matrix([x, y, z]))), " ",
 "ergibt sich die Berechnung wie folgt.</p><p>\\(",
 tex1(matrix([a, b, c])), tex1(transpose(matrix([x, y, z]))),
 "\\(=a \\times x + b \\times y + c \\times z \\)</p>",
 "<p>Auf diese Weise ergibt sich durch die Summe der Produkte entsprechender Komponenten eine einzelne Zahl.</p>",
 "<p>Als konkretes Beispiel ist für den Zeilenvektor ", tex1(matA0), " und den Spaltenvektor ", tex1(matB0), " das Produkt</p>",
 "<p style=\"text-align: center;\">", tex1(matA0B0S), "</p><p>.</p>",
 "<p>Versuchen wir nun einige Berechnungen.</p>")],
 ["ja", sconcat("<h2>行ベクトルと列ベクトルの積</h2>",
 "<p>行ベクトルと列ベクトルの積は, 対応する各要素の積をすべて足し合わせることで求められ, 1つの数値が得られます.</p>",
 "<p>例として、行ベクトル", tex1(matrix([a, b, c])), "と列ベクトル, ", tex1(transpose(matrix([x, y, z]))), "の場合, ",
 "計算は以下のようになります. </p><p>\\(",
 tex1(matrix([a, b, c])), tex1(transpose(matrix([x, y, z]))),
 "\\(=a \\times x + b \\times y + c \\times z \\)</p>",
 "<p>このように, 各要素同士の積の和を計算することで, 行ベクトルと列ベクトルの積は1つの数となります.</p>",
 "<p>具体例として, 行ベクトル", tex1(matA0), "と列ベクトル", tex1(matB0), "の場合, 積は</p>",
 "<p style=\"text-align: center;\">", tex1(matA0B0S), "</p><p>となります.</p>",
 "<p>では, いくつか計算してみましょう.</p>")]];
/* qtext:get_lang_string(qtextL, %_STACK_LANG); */
]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="html">
      <text>{@matA1@},{@matB1@}, {@matA1B1S@}/ {@matA2@},{@matB2@}, {@matA2B2S@}{@matA3@},{@matB3@}, {@matA3B3S@}</text>
    </questionnote>
    <questiondescription format="html">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> 
<span lang="ja" class="multilang">よくできました。正解です！</span>
<span lang="en" class="multilang">Well done! Spot on!</span>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span>
<span lang="ja" class="multilang">部分的に正解です。</span>
<span lang="en" class="multilang">Partially correct!</span>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span>
<span lang="ja" class="multilang">間違いです。</span>
<span lang="en" class="multilang">Incorrect!</span>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>(</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta1</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>allowempty,hideanswer</options>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>ta2</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>allowempty,hideanswer</options>
    </input>
    <input>
      <name>ans3</name>
      <type>numerical</type>
      <tans>ta3</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>allowempty,hideanswer</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>if ans1=EMPTYANSWER or ans2=EMPTYANSWER or ans3=EMPTYANSWER then empans:true else empans:false;
sc1:1/3;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>empty answer included?</description>
        <answertest>AlgEquiv</answertest>
        <sans>empans</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>すべての計算をしてください.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>(1)</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>sc1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>違います. 計算間違いでしょうか?
{@matA1@}{@matB1@}\(=\){@matA1B1S@}となります.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>(2)</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>sc1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text>正解です.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>違います. 計算間違いでしょうか?
{@matA2@}{@matB2@}\(=\){@matA2B2S@}となります.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>(3)</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>sc1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>違います. 計算間違いでしょうか?
{@matA3@}{@matB3@}\(=\){@matA3B3S@}となります.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>151014338</deployedseed>
    <deployedseed>1453718657</deployedseed>
    <deployedseed>1216702620</deployedseed>
    <deployedseed>2042663306</deployedseed>
    <deployedseed>378845994</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>管理者のインプットを想定したテストケースは満点です。</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-4-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>