<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 116214  -->
  <question type="stack">
    <name>
      <text>010.MatMapFindImageGeneral-A</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p><br></p>
<p>ある行列の定める1次写像\(f:{\mathbb K}^{@n@}\to{\mathbb K}^{@m@}\)がある.</p>
<p>\(f\)による{@v1@}の像は{@w1@}であり,{@v2@}の像は{@w2@}である.以下の問に答えよ</p>
<table>
    <tbody>
        <tr>
            <th>(1)</th>
            <td colspan="2">{@vv@}\(=\)[[input:ans1]][[validation:ans1]]{@v1@}\(+\)[[input:ans2]][[validation:ans2]]{@v2@}である.空欄を埋めよ.</td>
        </tr>
        <tr>
            <th>(2)</th>
            <td>{@vv@}の\(f\)による像を求めよ:</td>
            <td>[[input:ans3]][[validation:ans3]]</td>
        </tr>
    </tbody>
</table>
<script>
  var divs = document.getElementsByClassName('varmatrixinput');
  for(var i=0; i < divs.length; i++) { 
    divs[i].style.height = '120px';
  }
</script>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* 1次結合の1次写像による像 */
/* v002: 解答欄補正 (2023.4.19) */
%_nr:rand(3)+2;
%_nc:rand(2)+3;

/* m x n */
matA:genmatrix(lambda([i,j],rand([-2, -1, -1, 0, 1, 1, 1, 2])), %_nr, %_nc);
getRank2(n):=block([M,k],M:genmatrix(lambda([i,j], rand([-1, 0, 0, 1, 1, 1, 2])), %_nc, 2), 
   for k:0 while (k<100 and rank(M)<2) do 
     (M:genmatrix(lambda([i, j],rand([-1, 0, 0, 1, 1, 1, 2])), %_nc, 2)), return(M));
matV:getRank2(%_nc);
v1:col(matV,1);
v2:col(matV,2);
w1:matA.v1;
w2:matA.v2;
c1:rand([-2,-1,2,3,4]);
c2:rand([-3,-2,2,3,4]);
vv:c1*v1+c2*v2;
ww:matA.vv;
ta1:c1;
ta2:c2;
ta3:ww;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>({@v1@},{@v2@}),{@matA@},({@c1@},{@c2@})</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>よくできました。正解です!</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>惜しい！部分的に正解です。</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>残念，間違いです。</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>(</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>allowempty, hideanswer</options>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>allowempty, hideanswer</options>
    </input>
    <input>
      <name>ans3</name>
      <type>varmatrix</type>
      <tans>ta3</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>allowempty, hideanswer</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>(1) check</description>
        <answertest>AlgEquiv</answertest>
        <sans>[ans1, ans2]</sans>
        <tans>[ta1, ta2]</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.4</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>(1) 正解です.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>(1) が違います.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>(2) check</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.6</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>(2) 正解です.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>(2) 違います. (1)が正解のときは, 線形性を確認しよう. (1)が不正解のときは(1)をやり直そう.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1483871913</deployedseed>
    <deployedseed>173874659</deployedseed>
    <deployedseed>310235851</deployedseed>
    <deployedseed>1802587287</deployedseed>
    <deployedseed>1771070820</deployedseed>
    <deployedseed>215327868</deployedseed>
    <deployedseed>893616565</deployedseed>
    <deployedseed>816403906</deployedseed>
    <deployedseed>796342814</deployedseed>
    <deployedseed>1400208385</deployedseed>
    <deployedseed>736798390</deployedseed>
    <deployedseed>1309373768</deployedseed>
    <deployedseed>105350297</deployedseed>
    <deployedseed>722247737</deployedseed>
    <deployedseed>1479377536</deployedseed>
    <deployedseed>1173718084</deployedseed>
    <deployedseed>1681813989</deployedseed>
    <deployedseed>2020252106</deployedseed>
    <deployedseed>628726333</deployedseed>
    <deployedseed>1504146506</deployedseed>
    <deployedseed>291553251</deployedseed>
    <deployedseed>922130695</deployedseed>
    <deployedseed>1045274006</deployedseed>
    <deployedseed>479548436</deployedseed>
    <deployedseed>411910955</deployedseed>
    <deployedseed>494043652</deployedseed>
    <deployedseed>1619098363</deployedseed>
    <deployedseed>836483876</deployedseed>
    <deployedseed>238142225</deployedseed>
    <deployedseed>2001010100</deployedseed>
  </question>

</quiz>