<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 116209  -->
  <question type="stack">
    <name>
      <text>010.MatMapFindImageR3toR2-A</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p><br></p>
<div>\( \def\WR{\mathbb R}\def\retube#1#2{\begin{pmatrix}#1\\#2\end{pmatrix}}&nbsp; \)</div>
<p>線型写像 \( f:\WR^{@m@}\to \WR^{@n@} \) は, \({@WR@}^m\)の基本ベクトル {@e1@} を {@av1@} に, {@e2@} を {@av2@} に写す. このとき, \( {@Wv@}= \){@xv@} の \( f \) による像を答えよ.</p>
<table>
    <tbody>
        <tr>
            <th>\(f({@Wv@})=\)</th>
            <td>[[input:ans1]] [[validation:ans1]]</td>
        </tr>
    </tbody>
</table>
<script>
  var divs = document.getElementsByClassName('varmatrixinput');
  for(var i=0; i < divs.length; i++) { 
    divs[i].style.height = '120px';
  }
</script>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* 数ベクトル空間の間の線形写像, 基本ベクトルの像から計算 */
/* v003: n,m 汎用化対応 (2023.4.19) */
/* v002: パラメータ生成を整理, 問題文等全体調整 (2023.4.19) */
%_nr:2;%_nc:3;
matA:genmatrix(lambda([i, j],  rand([-3, -2, -1, -1, 0, 1, 1, 2, 3])), %_nr, %_nc);
j12:random_permutation(makelist(k, k, 1, %_nc));
j1:j12[1];j2:j12[2];
av1:col(matA, j1);
av2:col(matA, j2);
e1:genmatrix(lambda([i, j], if i=j1 then 1 else 0), %_nc, 1);
e2:genmatrix(lambda([i, j], if i=j2 then 1 else 0), %_nc, 1);
xv:genmatrix(lambda([i, j], if i=j1 then rand([-1, 1, 2]) elseif i=j2 then rand([-2, -1, 2]) else 0), %_nc, 1);
ta1:matA.xv;
texput(WR, "{\\mathbb R}");
texput(Wv, "{\\boldsymbol v}");
texput(We, "{\\boldsymbol e}");]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@matA@},{@xv@},{@Cans1@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>よくできました。正解です!</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>惜しい！部分的に正解です。</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>残念，間違いです。</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>(</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>varmatrix</type>
      <tans>ta1</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>[[BASIC-ALGEBRA]],[[BASIC-CALCULUS]],[[BASIC-MATRIX]]</forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>allowempty, hideanswer</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>chk0:is(matrix_size(ans1)=matrix_size(ta1));</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>chk0</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">答のベクトルのサイズが違います. 行列とベクトルの積による写像の定義を確認して下さい.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">残念ながら, 計算間違いのようです.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>675660902</deployedseed>
    <deployedseed>1146341449</deployedseed>
    <deployedseed>1982458929</deployedseed>
    <deployedseed>860035339</deployedseed>
    <deployedseed>392198790</deployedseed>
    <deployedseed>1955089701</deployedseed>
    <deployedseed>1579544871</deployedseed>
    <deployedseed>949952921</deployedseed>
    <deployedseed>751325844</deployedseed>
    <deployedseed>787096386</deployedseed>
    <deployedseed>2111735773</deployedseed>
    <deployedseed>197204324</deployedseed>
    <deployedseed>1671293354</deployedseed>
    <deployedseed>1497622458</deployedseed>
    <deployedseed>1835519714</deployedseed>
    <deployedseed>511597296</deployedseed>
    <deployedseed>1507466995</deployedseed>
    <deployedseed>1467021999</deployedseed>
    <deployedseed>1813673128</deployedseed>
    <deployedseed>716118845</deployedseed>
    <deployedseed>1482226788</deployedseed>
    <deployedseed>467196435</deployedseed>
    <deployedseed>517265260</deployedseed>
    <deployedseed>484806994</deployedseed>
    <deployedseed>1091495253</deployedseed>
    <deployedseed>2023958787</deployedseed>
    <deployedseed>769836450</deployedseed>
    <deployedseed>2018831704</deployedseed>
    <deployedseed>1545113524</deployedseed>
    <deployedseed>1256994972</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>管理者のインプットを想定したテストケースは満点です。</description>
      <testinput>
        <name>ans1</name>
        <value>Cans1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>